Electrophysiological dynamics of salience, default mode, and frontoparietal networks during episodic memory formation and recall: A multi-experiment iEEG replication

Author response:

The following is the authors’ response to the previous reviews.

Removing claims of causality: To avoid confusion, we have now removed claims of causality from our manuscript and also changed the title of the manuscript accordingly

"Electrophysiological dynamics of salience, default mode, and frontoparietal networks during episodic memory formation and recall: A multi-experiment iEEG replication".

Control analyses directly comparing AI and IFG: As per the reviewer’s suggestion, we have carried out additional control analyses by directly comparing the net inward/outward balance between the AI and the IFG. Our analysis revealed that the net outflow for the AI is significantly higher compared to the IFG during both encoding and recall phases, a pattern that was replicated across all four experiments.

These findings further highlight the unique role of the AI as a key hub in coordinating network interactions during episodic memory formation and retrieval, distinguishing it from a key anatomically adjacent prefrontal region implicated in cognitive control.

We have incorporated these results into the manuscript (see new Figure S6 and updated Results section).

Control analyses directly comparing task with resting state: As per the reviewer’s suggestion, we compared the AI's net outflow during task periods to resting state, finding significantly higher outflow during both encoding and recall across all experiments (ps < 0.05). These results provide further evidence for enhanced role of AI net directed information flow to the DMN and FPN during memory processing compared to the resting state.

We have incorporated these results into the manuscript (see new Figure S9 and updated Results section).

Control analysis using every region of the brain outside the considered networks: We appreciate the reviewer's suggestion to conduct additional control analyses. However, we have concerns about implementing this approach for several reasons:

(1) Hypothesis-driven research: Our study was designed based on a strong hypothesis derived from prior fMRI studies, which have consistently shown that the salience network (SN), anchored by the anterior insula (AI), plays a critical role in regulating the engagement and disengagement of the default mode network (DMN) and frontoparietal network (FPN) across diverse cognitive tasks.

(2) Risk of p-hacking: Running analyses on a large number of brain regions outside our networks of interest without a priori hypotheses could lead to p-hacking, a practice strongly criticized in the scientific community, including by eLife editors (Makin & Orban de Xivry, 2019). Such an approach could potentially yield spurious results and undermine the validity of our findings.

(3) Principled control region selection: Our choice of the inferior frontal gyrus (IFG) as a control region was hypothesis-driven, based on its: a) Anatomical adjacency to the AI b) Involvement in cognitive control functions, including response inhibition c) Frequent coactivation with the AI in fMRI studies.

(4) Robustness of current findings: Our PTE analysis involving the IFG, along with the additional control analyses requested by the reviewer (comparing the task-related net balance of the AI with the IFG and with resting state, see response to reviewer comment 2.1), strongly support a key role for the AI in orchestrating large-scale network dynamics during memory processes.

(5) Specificity of findings: The contrast between AI and IFG results demonstrates that our observed patterns are not general to all task-active regions but are specific to the AI's role in network coordination.

We believe that our current analyses, including the additional controls, provide a comprehensive and rigorous examination of the AI's role in memory-related network dynamics. Adding analyses of numerous additional regions without clear hypotheses could potentially dilute the focus and interpretability of our results.

However, we acknowledge the importance of considering broader network interactions. In future studies, we could explore the role of other key regions in a hypothesis-driven manner, potentially expanding our understanding of the complex interactions between multiple brain networks during memory processes.

These revisions, combined with our rigorous methodologies and comprehensive analyses, provide compelling support for the central claims of our manuscript. We believe these changes significantly enhance the scientific contribution of our work.

Our point-by-point responses to the reviewers' comments are provided below.

Reviewer 1:

(1.1) Because phase-transfer entropy is referenced as a "causal" analysis in this investigation (PTE), I believe it is important to highlight for readers recent discussions surrounding the description of "causal mechanisms" in neuroscience (see "Confusion about causation" section from Ross and Bassett, 2024, Nature Neuroscience). A large proportion of neuroscientists (myself included) use "causal" only to refer to a mechanism whose modulation or removal (with direct manipulation, such as by lesion or stimulation) is known to change or control a given outcome (such as a successful behavior). As Ross and Bassett highlight, it is debatable whether such mechanistic causality is captured by Granger "causality" (a.k.a. Granger prediction) or the parametric PTE, and imprecise use of "causation" may be confusing. The authors have defined in the revised Introduction what their definition of "causality" is within the context of this investigation.

We appreciate the reviewer's feedback in terms of the terminology used in our manuscript. To avoid confusion, we have now removed claims of causality from our manuscript and also changed the title of the manuscript accordingly.

Reviewer 2:

(2.1) Clarifying the new control analyses. The authors have been responsive to our feedback and implemented several new analyses. The use of a pre-task baseline period and a control brain region (IFG) definitively help to contextualize their results, and the findings shown in the revision do suggest that (1) relative to a pre-task baseline, directed interactions from the AI are stronger and (2) relative to a nearby region, the IFG, the AI exhibits greater outward-directed influence.

However, it is difficult to draw strong quantitative conclusions from the analyses as presented, because they do not directly statistically contrast the effect in question (directed interactions with the FPN and DMN) between two conditions (e.g. during baseline vs. during memory encoding/retrieval). As I understand it, in their main figures the authors ask, "Is there statistically greater influence from the AI to the DMN/FPN in one direction versus another?" And in the AI they show greater "outward" PTE than "inward" PTE from other networks during encoding/retrieval. The balance of directed information favors an outward influence from the AI to DMN/FPN.

But in their new analyses, they simply show that the degree of "outward" PTE is greater during task relative to baseline in (almost) all tasks. I believe a more appropriately matched analysis would be to quantify the inward/outward balance during task states, quantify the inward/outward balance during rest states, and then directly statistically compare the two. It could be that the relative balance of directed information flow is nonsignificantly changed between task and rest states, which would be important to know.

We thank the reviewer for this suggestion. We have now run additional analysis by directly comparing the inward/outward balance during the task versus the rest states. To calculate the net inward/outward balance, we calculated the net outflow as the difference between the total outgoing information and total incoming information (PTE(out)–PTE(in)). This analysis revealed that net outflow during task periods is significantly higher compared to rest, during both encoding and recall, and across the four experiments (ps < 0.05). These results provide further evidence for enhanced role of AI net directed information flow to the DMN and FPN during memory processing compared to the resting state. These new results have now been included in the revised manuscript (page 12).

Likewise, a similar principle applies to their IFG analysis. They show that the IFG tends to have an "inward" balance of influence from the DMN/FPN (the opposite of the AIs effect), but this does not directly answer whether the AI occupies a statistically unique position in terms of the magnitude of its influence on other regions. More appropriate, as I suggest above, would be to quantify the relative balance inward/outward influence, both for the IFG and the AI, and then directly compare those two quantities. (Given the inversion of the direction of effect, this is likely to be a significant result, but I think it deserves a careful approach regardless.)

We appreciate the reviewer's suggestion. As per the reviewer’s suggestion, we directly compared the net inward/outward balance between the AI and the IFG. Specifically, we compared the net outflow (PTE(out)–PTE(in)) for the AI with the IFG. This analysis revealed that the net outflow for the AI is significantly higher compared to the IFG during both encoding and recall, and across the four experiments. These findings further highlight a key role for the AI in orchestrating large-scale network dynamics during memory processes. The AI's pattern of directed information flow stands in contrast to that of the IFG, despite their anatomical proximity and shared involvement in cognitive control processes. This dissociation underscores the specificity of the AI's function in coordinating network interactions during memory formation and retrieval. These new results have now been included in our revised manuscript (page 11).

(2.2) Consider additional control regions. The authors justify their choice of IFG as a control region very well. In my original comments, I perhaps should have been more clear that the most compelling control analyses here would be to subject every region of the brain outside these networks (with good coverage) to the same analysis, quantify the degree of inward/outward balance, and then see how the magnitude of the AI effect stacks up against all possible other options. If the assertion is that the AI plays a uniquely important role in these memory processes, showing how its influence stacks up against all possible "competitors" would be a very compelling demonstration of their argument.

We thank the reviewer for this suggestion. However, please note that running a large number of random analysis by including a large number of brain regions (every region of the brain outside these networks) and comparing their dynamics to the AI without a hypothesis or solid principle amounts to p-hacking, which has been previously strongly criticized by the eLife editors (Makin & Orban de Xivry, 2019). Our study was strongly driven by a solid hypothesis based on prior fMRI studies that have shown that the SN, anchored by the anterior insula (AI), plays a critical role in regulating the engagement and disengagement of the DMN and FPN across diverse cognitive tasks (Bressler & Menon, 2010; Cai et al., 2016; Cai, Ryali, Pasumarthy, Talasila, & Menon, 2021; Chen, Cai, Ryali, Supekar, & Menon, 2016; Kronemer et al., 2022; Raichle et al., 2001; Seeley et al., 2007; Sridharan, Levitin, & Menon, 2008). Moreover, our selection of the IFG as a control region for comparison was also very strongly hypothesis driven, due to its anatomical adjacency to the AI, its involvement in a wide range of cognitive control functions including response inhibition (Cai, Ryali, Chen, Li, & Menon, 2014), and its frequent co-activation with the AI in fMRI studies. Furthermore, the IFG has been associated with controlled retrieval of memory (Badre, Poldrack, Paré-Blagoev, Insler, & Wagner, 2005; Badre & Wagner, 2007; Wagner, Paré-Blagoev, Clark, & Poldrack, 2001), making it a compelling region for comparison. Our findings related to the PTE analysis involving the IFG and also the additional control analyses requested by the reviewer (directly comparing the task-related net balance of the AI with the IFG and also to resting state, please see response to reviewer comment 2.1) strongly highlight a key role of the AI in orchestrating large-scale network dynamics during memory processes.

We believe that our current analyses, including the additional controls, provide a comprehensive and rigorous examination of the AI's role in memory-related network dynamics. Adding analyses of numerous additional regions without clear hypotheses could potentially dilute the focus and interpretability of our results.

However, we acknowledge the importance of considering broader network interactions. In future studies, we could explore the role of other key regions in a hypothesis-driven manner, potentially expanding our understanding of the complex interactions between multiple brain networks during memory processes.

(2.3) Reporting of successful vs. unsuccessful memory results. I apologize if I was not clear in my original comment (2.7, pg. 13 of the response document) regarding successful vs. unsuccessful memory. The fact that no significant difference was found in PTE between successful/unsuccessful memory is a very important finding that adds valuable context to the rest of the manuscript. I believe it deserves a figure, at least in the Supplement, so that readers can visualize the extent of the effect in successful/unsuccessful trials. This is especially important now that the manuscript has been reframed to focus more directly on claims regarding episodic memory processing; if that is indeed the focus, and their central analysis does not show a significant effect conditionalized on the success of memory encoding/retrieval, it is important that readers can see these data directly.

As per the reviewer’s suggestion, we have now included a Figure related to the results for the successful versus unsuccessful comparison in the Supplementary materials of the revised manuscript (Figures S10, S11).

(2.4) Claims regarding causal relationships in the brain. I understand that the authors have defined "causal" in a specific way in the context of their manuscript; I do believe that as a matter of clear and transparent scientific communication, the authors nonetheless bear a responsibility to appreciate how this word may be erroneously interpreted/overinterpreted and I would urge further review of the manuscript to tone down claims of causality. Reflective of this, I was very surprised that even as both reviewers remarked on the need to use the word "causal" with extreme caution, the authors added it to the title in their revised manuscript.

We thank the reviewer for this suggestion. To avoid confusion, we have now removed claims of causality from our manuscript and also changed the title of the manuscript accordingly.

References

Badre, D., Poldrack, R. A., Paré-Blagoev, E. J., Insler, R. Z., & Wagner, A. D. (2005). Dissociable controlled retrieval and generalized selection mechanisms in ventrolateral prefrontal cortex. Neuron, 47(6), 907-918. doi:10.1016/j.neuron.2005.07.023

Badre, D., & Wagner, A. D. (2007). Left ventrolateral prefrontal cortex and the cognitive control of memory. Neuropsychologia, 45(13), 2883-2901. doi:10.1016/j.neuropsychologia.2007.06.015

Bressler, S. L., & Menon, V. (2010). Large-scale brain networks in cognition: emerging methods and principles. Trends in Cognitive Sciences, 14(6), 277-290. doi:10.1016/j.tics.2010.04.004

Cai, W., Chen, T., Ryali, S., Kochalka, J., Li, C. S., & Menon, V. (2016). Causal Interactions Within a Frontal-Cingulate-Parietal Network During Cognitive Control: Convergent Evidence from a Multisite-Multitask Investigation. Cereb Cortex, 26(5), 2140-2153. doi:10.1093/cercor/bhv046

Cai, W., Ryali, S., Chen, T., Li, C. S., & Menon, V. (2014). Dissociable roles of right inferior frontal cortex and anterior insula in inhibitory control: evidence from intrinsic and taskrelated functional parcellation, connectivity, and response profile analyses across multiple datasets. J Neurosci, 34(44), 14652-14667. doi:10.1523/jneurosci.3048-14.2014

Cai, W., Ryali, S., Pasumarthy, R., Talasila, V., & Menon, V. (2021). Dynamic causal brain circuits during working memory and their functional controllability. Nat Commun, 12(1), 3314. doi:10.1038/s41467-021-23509-x

Chen, T., Cai, W., Ryali, S., Supekar, K., & Menon, V. (2016). Distinct Global Brain Dynamics and Spatiotemporal Organization of the Salience Network. PLOS Biology, 14(6), e1002469. doi:10.1371/journal.pbio.1002469

Kronemer, S. I., Aksen, M., Ding, J. Z., Ryu, J. H., Xin, Q., Ding, Z., . . . Blumenfeld, H. (2022). Human visual consciousness involves large scale cortical and subcortical networks independent of task report and eye movement activity. Nat Commun, 13(1), 7342. doi:10.1038/s41467-022-35117-4

Makin, T. R., & Orban de Xivry, J. J. (2019). Ten common statistical mistakes to watch out for when writing or reviewing a manuscript. Elife, 8. doi:10.7554/eLife.48175

Raichle, M. E., MacLeod, A. M., Snyder, A. Z., Powers, W. J., Gusnard, D. A., & Shulman, G. L. (2001). A default mode of brain function. Proc Natl Acad Sci U S A, 98(2), 676-682. doi:10.1073/pnas.98.2.676

Seeley, W. W., Menon, V., Schatzberg, A. F., Keller, J., Glover, G. H., Kenna, H., . . . Greicius, M. D. (2007). Dissociable Intrinsic Connectivity Networks for Salience Processing and Executive Control. Journal of Neuroscience, 27(9), 2349-2356. doi:10.1523/JNEUROSCI.5587-06.2007

Sridharan, D., Levitin, D. J., & Menon, V. (2008). A critical role for the right fronto-insular cortex in switching between central-executive and default-mode networks. Proceedings of the National Academy of Sciences, 105(34), 12569-12574. doi:10.1073/pnas.0800005105

Wagner, A. D., Paré-Blagoev, E. J., Clark, J., & Poldrack, R. A. (2001). Recovering meaning: left prefrontal cortex guides controlled semantic retrieval. Neuron, 31(2), 329-338. doi:10.1016/s0896-6273(01)00359-2

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Dynamic interactions between large-scale brain networks underpin human cognitive processes, but their electrophysiological mechanisms remain elusive. The triple network model, encompassing the salience (SN), default mode (DMN), and frontoparietal (FPN) networks, provides a framework for understanding these interactions. To unravel the electrophysiological mechanisms underlying these network interactions, we analyzed intracranial EEG recordings from 177 participants across four diverse episodic memory experiments, each involving encoding as well as recall phases. Phase transfer entropy analysis revealed consistently higher directed information flow from the anterior insula, a key SN node, to both DMN and FPN nodes. This causal influence was significantly stronger during memory tasks compared to resting-state, highlighting the anterior insula’s task-specific role in coordinating large-scale network interactions. This pattern persisted across externally-driven memory encoding and internally-governed free recall. We also observed task-specific suppression of high-gamma power in the posterior cingulate cortex/precuneus node of the DMN during memory encoding, but not recall. Crucially, these results were robustly replicated across all four experiments spanning verbal and spatial memory domains with high Bayes replication factors. These findings significantly advance our understanding of how coordinated neural network interactions support memory processes. They highlight the anterior insula’s critical role in orchestrating large-scale brain network dynamics during both memory encoding and retrieval. By elucidating the electrophysiological basis of triple network interactions in episodic memory, our results provide insights into neural circuit dynamics underlying memory function and offer a framework for investigating network disruptions in neurological and psychiatric disorders affecting memory.

Revised version: Dynamic interactions between large-scale brain networks underpin human cognitive processes, but their electrophysiological mechanisms remain elusive. The triple network model, encompassing the salience (SN), default mode (DMN), and frontoparietal (FPN) networks, provides a framework for understanding these interactions. We analyzed intracranial EEG recordings from 177 participants across four diverse episodic memory experiments, each involving encoding as well as recall phases. Phase transfer entropy analysis revealed consistently higher directed information flow from the anterior insula (AI), a key SN node, to both DMN and FPN nodes. This directed influence was significantly stronger during memory tasks compared to resting-state, highlighting the AI’s task-specific role in coordinating large-scale network interactions. This pattern persisted across externally-driven memory encoding and internally-governed free recall. Control analyses using the inferior frontal gyrus (IFG) showed an inverse pattern, with DMN and FPN exerting higher influence on IFG, underscoring the AI’s unique role. We observed task-specific suppression of high-gamma power in the posterior cingulate cortex/precuneus node of the DMN during memory encoding, but not recall. Crucially, these results were replicated across all four experiments spanning verbal and spatial memory domains with high Bayes replication factors. Our findings advance understanding of how coordinated neural network interactions support memory processes, highlighting the AI’s critical role in orchestrating large-scale brain network dynamics during both memory encoding and retrieval. By elucidating the electrophysiological basis of triple network interactions in episodic memory, our study provides insights into neural circuit dynamics underlying memory function and offer a framework for investigating network disruptions in memory-related disorders.

Introduction

Dynamic interactions between large-scale brain networks are thought to underpin human cognitive processes, but the electrophysiological dynamics that underlie these interactions remain elusive. The triple network model, which includes the salience (SN), default mode (DMN), and frontoparietal (FPN) networks, offers a fundamental framework for understanding these complex interactions (Cai, Ryali, Pasumarthy, Talasila, & Menon, 2021; Menon, 2011, 2023). These networks collaboratively manage tasks that require significant stimulus-driven and stimulus-independent attentional control, highlighting the integrated nature of brain function. Building on Mesulam’s (Mesulam, 1990) theory that all cognitive and memory systems operate within a complex architecture of interconnected brain regions, the triple network model articulates how these networks facilitate demanding cognitive tasks. However, despite the model’s broad influence, the specific electrophysiological mechanisms that support these interactions during cognition remain poorly understood.

Episodic memory, the cognitive process of encoding, storing, and retrieving personally experienced events, is essential for a variety of complex cognitive functions and everyday activities (Dickerson & Eichenbaum, 2010; Düzel, Penny, & Burgess, 2010; Moscovitch, Cabeza, Winocur, & Nadel, 2016; Ranganath & Ritchey, 2012; Rugg & Vilberg, 2013; Rutishauser, Reddy, Mormann, & Sarnthein, 2021; Yonelinas, Ranganath, Ekstrom, & Wiltgen, 2019). Influential theoretical models of human memory posit a key role for control processes in regulating hierarchical processes associated with episodic memory formation (Andermane, Joensen, & Horner, 2021; Atkinson & Shiffrin, 1968; Bastos et al., 2012; Kumaran & McClelland, 2012; Tulving, 2002). Crucially, the formation of episodic memories relies on the intricate interplay between external stimulus-driven processes during encoding and internal recall processes during retrieval (Buckner & DiNicola, 2019; Fornito, Harrison, Zalesky, & Simons, 2012; Mesulam, 1990), making it an ideal cognitive process to investigate the triple network model’s broader applicability and its underlying neurophysiological mechanisms. Elucidating these mechanisms is crucial not only for understanding basic brain functions but also for addressing neuropsychological disorders where these mechanisms may be disrupted (Li et al., 2019).

Each network in the triple network model plays a unique and critical role in regulating human cognition (Menon, 2023). The SN, anchored by the anterior insula (AI), identifies and filters salient stimuli, helping individuals focus on goal-relevant aspects of their environment (Menon & Uddin, 2010). In contrast, the DMN is typically engaged during internally focused cognitive processes and is implicated in retrieval of past events and experiences (Buckner, Andrews-Hanna, & Schacter, 2008; Fox & Raichle, 2007; Fox et al., 2005; Greicius et al., 2008; Greicius & Menon, 2004; Laufs et al., 2003; Raichle, 2015; Raichle et al., 2001; Smallwood et al., 2021). The FPN is involved in the maintenance and manipulation of information within working memory and exerts top-down attentional control to regulate memory formation (Badre, Poldrack, Paré-Blagoev, Insler, & Wagner, 2005; Badre & Wagner, 2007; Helfrich & Knight, 2016; Jin, Olk, & Hilgetag, 2010; Simons & Spiers, 2003; Uncapher & Wagner, 2009; Wagner, Paré-Blagoev, Clark, & Poldrack, 2001; Wagner, Shannon, Kahn, & Buckner, 2005).

Earlier version: Central to the functionality of this model is the AI, a pivotal node within the SN. Functional brain imaging studies have revealed the SN’s critical role in regulating the engagement and disengagement of the DMN and FPN across diverse cognitive tasks (Bressler & Menon, 2010; Cai et al., 2016; Cai et al., 2021; Chen, Cai, Ryali, Supekar, & Menon, 2016; Kronemer et al., 2022; Raichle et al., 2001; Seeley et al., 2007; Sridharan, Levitin, & Menon, 2008). The AI dynamically detects and filters task-relevant information, facilitating rapid and efficient switching between the DMN and FPN in response to shifting task demands (Menon, 2015a). However, how this process operates at the neurophysiological level remains unknown, underlining a significant gap in our understanding of causal network dynamics in memory formation.

Revised version: Central to the functionality of this model is the AI, a pivotal node within the SN. Functional brain imaging studies have revealed the SN’s critical role in regulating the engagement and disengagement of the DMN and FPN across diverse cognitive tasks (Bressler & Menon, 2010; Cai et al., 2016; Cai et al., 2021; Chen, Cai, Ryali, Supekar, & Menon, 2016; Kronemer et al., 2022; Raichle et al., 2001; Seeley et al., 2007; Sridharan, Levitin, & Menon, 2008). The AI dynamically detects and filters task-relevant information, facilitating rapid and efficient switching between the DMN and FPN in response to shifting task demands (Menon, 2015a). However, how this process operates at the neurophysiological level remains unknown, underlining a significant gap in our understanding of directed network dynamics in memory formation.

While the tripartite network has been most extensively studied in the context of cognitive tasks requiring explicit cognitive control, growing evidence suggests its relevance to episodic memory as a domain-general control system. Brain imaging studies in both healthy individuals and clinical populations provide growing evidence for the involvement of the tripartite network in memory processes. In healthy adults, Sestieri et al. found that the SN exhibited sustained activity across all phases of both episodic memory search and perceptual tasks (Sestieri, Corbetta, Spadone, Romani, & Shulman, 2014). The SN was consistently activated across all task phases, from initiation to response, indicating its broad involvement in memory processes. Importantly, the SN demonstrated flexible functional connectivity, linking with the DMN during memory search and dorsal attention network during perceptual search. These findings point to the SN’s involvement in dynamically coordinating large-scale brain networks during episodic memory processes, supporting its characterization as a versatile, domain-general control network that adapts its connectivity patterns to meet diverse cognitive demands.

Further supporting this view, Vatansever et al. demonstrated shared neural processes, centered on the AI, supporting the controlled retrieval of both semantic and episodic memories (Vatansever, Smallwood, & Jefferies, 2021). They identified a common cluster of cortical activity centered on the AI and adjoining inferior frontal gyrus for the retrieval of both weakly-associated semantic and weakly-encoded episodic memory traces. Moreover, they found that reduced functional interaction between this cluster and the ventromedial prefrontal cortex, a key node of the DMN, was associated with better performance across both memory types. Higher pre-stimulus activity in the SN was associated with increased activity in temporal regions linked to encoding and reduced activity in regions associated with retrieval and self-referential processing (Cohen et al., 2020). This suggests that the SN may regulate memory by enhancing encoding and reducing interference from competing memory processes. Together, these findings not only reinforce the domain-general role of the SN in memory processes but also highlight the importance of investigating interactions between the tripartite network components during memory tasks.

Clinical studies have also underscored aspects of the tripartite network in memory function. Le Berre et al. found that disrupted insula connectivity was associated with unawareness of memory impairments in non-Korsakoff’s syndrome alcoholism, highlighting the crucial role of the right insula in memory functioning (Le Berre et al., 2017). Additionally, alcoholics showed weaker connectivity between the right insula and the dorsal anterior cingulate cortex nodes of the SN, and stronger connectivity between the right insula and ventromedial prefrontal cortex, a key node of the DMN. Importantly, alcoholics who failed to desynchronize insula-ventromedial prefrontal cortex activity demonstrated greater overestimation of their memory predictions and poorer recognition performance. Similarly, Xie et al. demonstrated that disrupted intrinsic connectivity of insula networks was associated with episodic memory deficits in patients with amnestic mild cognitive impairment (Xie et al., 2012). These studies suggest that disrupted insula connectivity may underlie the lack of awareness of memory impairments, and highlights the crucial role of the SN in memory functioning.

Earlier version: Despite these advances, the electrophysiological basis and dynamic interactions of these networks during memory formation and retrieval remain poorly understood. Our understanding of dynamic network interactions during human cognition is primarily informed by fMRI studies, which are limited by their temporal resolution. This constraint impedes our understanding of real-time, millisecond-scale neural dynamics and underscores the need to explore network interactions at time scales more pertinent to neural circuit dynamics. However, the difficulties involved in acquiring human electrophysiological data from multiple brain regions have made it challenging to elucidate the precise neural mechanisms underlying the functioning of large-scale networks. These challenges obscure our understanding of the dynamic temporal properties and causal interactions between the AI and other large-scale distributed networks during memory formation.

Revised version: Despite these advances, the electrophysiological basis and dynamic interactions of these networks during memory formation and retrieval remain poorly understood. Our understanding of dynamic network interactions during human cognition is primarily informed by fMRI studies, which are limited by their temporal resolution. This constraint impedes our understanding of real-time, millisecond-scale neural dynamics and underscores the need to explore network interactions at time scales more pertinent to neural circuit dynamics. However, the difficulties involved in acquiring human electrophysiological data from multiple brain regions have made it challenging to elucidate the precise neural mechanisms underlying the functioning of large-scale networks. These challenges obscure our understanding of the dynamic temporal properties and directed interactions between the AI and other large-scale distributed networks during memory formation.

To address these gaps, we leveraged intracranial EEG (iEEG) data acquired during multiple memory experiments from the University of Pennsylvania Restoring Active Memory (UPENN-RAM) study (Solomon et al., 2019). This dataset provides an unprecedented opportunity to probe the electrophysiological dynamics of triple network interactions during both episodic memory encoding and recall, with depth recordings from 177 participants across multiple memory experiments. The UPENN-RAM dataset includes electrodes in the AI, the posterior cingulate cortex (PCC)/precuneus and medial prefrontal cortex (mPFC) nodes of the DMN, and the dorsal posterior parietal cortex (dPPC) and middle frontal gyrus (MFG) nodes of the FPN. By examining four diverse episodic memory tasks spanning verbal and spatial domains, we aimed to elucidate the neurophysiological underpinnings of the AI’s dynamic network interactions with the DMN and FPN and assess the consistency of these interactions across tasks and stages of memory formation.

We investigated four episodic memory experiments spanning both verbal and spatial domains. The first experiment was a verbal free recall memory task (VFR) in which participants were presented with a sequence of words during the encoding period and asked to remember them for subsequent verbal recall. The second was a categorized verbal free recall task (CATVFR) in which participants were presented with a sequence of categorized words during the encoding period and asked to remember them for subsequent verbal recall. The third involved a paired associates learning verbal cued recall task (PALVCR) in which participants were presented with a sequence of word-pairs during the encoding period and asked to remember them for subsequent verbal cued recall. The fourth was a water maze spatial memory task (WMSM) in which participants were shown objects in various locations during the encoding periods and asked to retrieve the location of the objects during a subsequent recall period. This comprehensive approach afforded a rare opportunity in an iEEG setting to examine network interactions between the AI and the DMN and FPN nodes during both encoding and recall phases across multiple memory domains.

Earlier version: A crucial test of the triple network model is whether the AI exerts a strong directed causal influence on the DMN and FPN. The AI is consistently engaged during attentional tasks and, dynamic causal modeling of fMRI data suggests that it exerts strong causal influences on the DMN and FPN in these contexts (Cai et al., 2016; Cai et al., 2021; Chen et al., 2016; Sridharan et al., 2008; Wen, Liu, Yao, & Ding, 2013). However, it remains unknown whether the AI plays a causal role during both memory encoding and recall and whether such influences have a neurophysiological basis. To investigate the directionality of information flow between neural signals in the AI and DMN and FPN, we employed phase transfer entropy (PTE), a robust and powerful measure for characterizing information flow between brain regions based on phase coupling (Hillebrand et al., 2016; Lobier, Siebenhühner, Palva, & Matias, 2014; Wang et al., 2017). Crucially, it captures linear and nonlinear intermittent and nonstationary causal dynamics in iEEG data (Hillebrand et al., 2016; Lobier et al., 2014; Menon et al., 1996). We hypothesized that the AI would exert higher directed influence on the DMN and FPN than the reverse. We operationalize causality as follows: a brain region has a causal influence on a target if knowing the past history of temporal signals in both regions improves the ability to predict the target’s signal in comparison to knowing only the target’s past, as defined in earlier studies (Granger, 1969; Lobier et al., 2014).

Revised version: A crucial test of the triple network model is whether the AI exerts a strong directed directed influence on the DMN and FPN. The AI is consistently engaged during attentional tasks and, dynamic causal modeling of fMRI data suggests that it exerts strong causal influences on the DMN and FPN in these contexts (Cai et al., 2016; Cai et al., 2021; Chen et al., 2016; Sridharan et al., 2008; Wen, Liu, Yao, & Ding, 2013). However, it remains unknown whether the AI plays a causal role during both memory encoding and recall and whether such influences have a neurophysiological basis. To investigate the directionality of information flow between neural signals in the AI and DMN and FPN, we employed phase transfer entropy (PTE), a robust and powerful measure for characterizing information flow between brain regions based on phase coupling (Hillebrand et al., 2016; Lobier, Siebenhühner, Palva, & Matias, 2014; Wang et al., 2017). Crucially, it captures linear and nonlinear intermittent and nonstationary dynamics in iEEG data (Hillebrand et al., 2016; Lobier et al., 2014; Menon et al., 1996). We hypothesized that the AI would exert higher directed influence on the DMN and FPN than the reverse.

To further enhance our understanding of the dynamic activations within the three networks during episodic memory formation, we determined whether high-gamma band power in the AI, DMN, and FPN nodes depends on the phase of memory formation. Memory encoding, driven primarily by external stimulation, might invoke different neural responses compared to memory recall, which is more internally driven (Andrews-Hanna, 2012; Buckner et al., 2008). We hypothesized that DMN power would be suppressed during memory encoding as it is primarily driven by external stimuli, whereas an opposite pattern would be observed during memory recall which is more internally driven. Based on the distinct functions of the DMN and FPN— internally-oriented cognition and adaptive external response —we expected to observe differential modulations during encoding and recall phases. By testing these hypotheses, we aimed to provide a more detailed understanding of the dynamic role of triple network interactions in episodic memory formation, offering insights into the temporal dynamics and directed interactions within these large-scale cognitive networks.

Our final objective was to investigate the replicability of our findings across multiple episodic memory domains involving both verbal and spatial materials. Reproducing findings across experiments is a significant challenge in neuroscience, particularly in invasive iEEG studies where data sharing and sample sizes have been notable limitations. There have been few previous replicated findings from human iEEG studies across multiple task domains. Quantitatively rigorous measures are needed to address the reproducibility crisis in human iEEG studies. We used Bayesian analysis to quantify the degree of replicability (Ly, Etz, Marsman, & Wagenmakers, 2019; Verhagen & Wagenmakers, 2014). Bayes factors (BFs) are a powerful tool for evaluating evidence for replicability of findings across tasks and for determining the strength of evidence for the null hypothesis (Verhagen & Wagenmakers, 2014). Briefly, the replication BF is the ratio of marginal likelihood of the replication data, given the posterior distribution estimated from the original data, and the marginal likelihood for the replication data under the null hypothesis of no effect (Ly et al., 2019).

In summary, our study aims to elucidate the neurophysiological basis of the interactions between large-scale cognitive networks by leveraging a unique dataset of iEEG recordings across multiple memory experiments. By examining directed information flow, high-gamma band power modulation, and replicability across verbal and spatial memory domains, we sought to advance our understanding of the neural mechanisms underpinning human episodic memory. Our findings shed light on how the brain effectively integrates information from distinct networks to support memory formation, and cognition more broadly.

Results

AI response compared to PCC/precuneus during encoding and recall in the VFR task

We first examined neuronal activity in the AI and the PCC/precuneus and tested whether activity in the PCC/precuneus is suppressed compared to activity in the AI. Previous studies have suggested that power in the high-gamma band (80-160 Hz) is correlated with fMRI BOLD signals (Hermes, Nguyen, & Winawer, 2017; Hutchison, Hashemi, Gati, Menon, & Everling, 2015; Lakatos, Gross, & Thut, 2019; Leopold, Murayama, & Logothetis, 2003; Mantini, Perrucci, Del Gratta, Romani, & Corbetta, 2007; Schölvinck, Maier, Ye, Duyn, & Leopold, 2010), and is thought to reflect local neuronal activity (Canolty & Knight, 2010). Therefore, we compared high-gamma band power (see Methods for details) in the AI and PCC/precuneus electrodes during both encoding and recall and across the four episodic memory tasks. Briefly, in the VFR task, participants were presented with a sequence of words and asked to remember them for subsequent recall (Methods, Tables S1, S2a, S3a, Figures 1a, 2).

Added in the revised version: Encoding Compared to the AI, high-gamma power in PCC/precuneus was suppressed during almost the entire window 110 – 1600 msec during memory encoding (ps < 0.05, Figure 3a).

Earlier version: Task design of the encoding and recall periods of the memory experiments, and iEEG recording sites in AI, with DMN and FPN nodes.

Revised version: Task design of the encoding and recall periods of the memory experiments, and iEEG recording sites in AI, with DMN and FPN nodes. (a) Experiment 1, Verbal free recall (VFR):

Earlier version: (a) Experiment 1, Verbal free recall (VFR): (i) Task design of memory encoding and recall periods of the verbal free recall experiment (see Methods for details). Participants were first presented with a list of words in the encoding block and asked to recall as many as possible from the original list after a short delay. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the verbal free recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 9%, 8%, 19%, 32%, and 32% respectively, in the VFR experiment. (b) Experiment 2, Categorized verbal free recall (CATVFR): (i) Task design of memory encoding and recall periods of the categorized verbal free recall experiment (see Methods for details). Participants were presented with a list of words with consecutive pairs of words from a specific category (for example, JEANS-COAT, GRAPE-PEACH, etc.) in the encoding block and subsequently asked to recall as many as possible from the original list after a short delay. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the categorized verbal free recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 10%, 7%, 11%, 35%, and 37% respectively, in the CATVFR experiment. (c) Experiment 3, Paired associates learning verbal cued recall (PALVCR): (i) Task design of memory encoding and recall periods of the paired associates learning verbal cued recall experiment (see Methods for details). Participants were first presented with a list of 6 word-pairs in the encoding block and after a short post-encoding delay, participants were shown a specific word-cue and asked to verbally recall the cued word from memory. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the paired associates learning verbal cued recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 14%, 5%, 13%, 33%, and 35% respectively, in the PALVCR experiment. (d) Experiment 4, Water maze spatial memory (WMSM): (i) Task design of memory encoding and recall periods of the water maze spatial memory experiment (see Methods for details). Participants were shown objects in various locations during the encoding period and asked to retrieve the location of the objects during the recall period. Modified from Jacobs et. al. (2018) with permission. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the water maze spatial memory experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 10%, 15%, 13%, 38%, and 24% respectively, in the WMSM experiment. Overall, proportion of electrodes for VFR, CATVFR, PALVCR, and WMSM experiments were 43%, 27%, 15%, and 15% respectively. AI: anterior insula, PCC: posterior cingulate cortex, Pr: precuneus, mPFC: medial prefrontal cortex, dPPC: dorsal posterior parietal cortex, MFG: middle frontal gyrus.

Revised version: (i) Task design of memory encoding and recall periods of the verbal free recall experiment (see Methods for details). Participants were first presented with a list of words in the encoding block and asked to recall as many as possible from the original list after a short delay. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the verbal free recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 9%, 8%, 19%, 32%, and 32% respectively, in the VFR experiment. (b) Experiment 2, Categorized verbal free recall (CATVFR): (i) Task design of memory encoding and recall periods of the categorized verbal free recall experiment (see Methods for details). Participants were presented with a list of words with consecutive pairs of words from a specific category (for example, JEANS-COAT, GRAPE-PEACH, etc.) in the encoding block and subsequently asked to recall as many as possible from the original list after a short delay. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the categorized verbal free recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 10%, 7%, 11%, 35%, and 37% respectively, in the CATVFR experiment. (c) Experiment 3, Paired associates learning verbal cued recall (PALVCR): (i) Task design of memory encoding and recall periods of the paired associates learning verbal cued recall experiment (see Methods for details). Participants were first presented with a list of 6 word-pairs in the encoding block and after a short post-encoding delay, participants were shown a specific word-cue and asked to verbally recall the cued word from memory. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the paired associates learning verbal cued recall experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 14%, 5%, 13%, 33%, and 35% respectively, in the PALVCR experiment. (d) Experiment 4, Water maze spatial memory (WMSM): (i) Task design of memory encoding and recall periods of the water maze spatial memory experiment (see Methods for details). Participants were shown objects in various locations during the encoding period and asked to retrieve the location of the objects during the recall period. Modified from Jacobs et. al. (2018) with permission. (ii) Electrode locations for AI with DMN nodes (top panel) and AI with FPN nodes (bottom panel), in the water maze spatial memory experiment. Proportion of electrodes for AI, PCC/Pr, mPFC, dPPC, and MFG were 10%, 15%, 13%, 38%, and 24% respectively, in the WMSM experiment. Overall, proportion of electrodes for VFR, CATVFR, PALVCR, and WMSM experiments were 43%, 27%, 15%, and 15% respectively. AI: anterior insula, PCC: posterior cingulate cortex, Pr: precuneus, mPFC: medial prefrontal cortex, dPPC: dorsal posterior parietal cortex, MFG: middle frontal gyrus.

Anterior insula electrode locations (red) visualized on insular regions based on the atlas by Faillenot and colleagues (Faillenot, Heckemann, Frot, & Hammers, 2017).

Anterior insula is shown in blue, and posterior insula mask is shown in green (see Methods for details). This atlas is based on probabilistic analysis of the anatomy of the insula with demarcations of the AI based on three short dorsal gyri and the PI which encompasses two long and ventral gyri.

Removed in the revised version: Encoding Compared to the AI, high-gamma power in PCC/precuneus was suppressed during almost the entire window 110 – 1600 msec during memory encoding (ps < 0.05, Figure 3a).

iEEG evoked response, quantified using high-gamma (HG) power, for AI (red) and PCC/precuneus (blue) during (a) VFR, (b) CATVFR, (c) PALVCR, and (d) WMSM experiments.

Green horizontal lines denote greater high-gamma power for AI compared to PCC/precuneus (ps < 0.05). Red horizontal lines denote increase of AI response compared to the resting baseline during the encoding and recall periods (ps < 0.05). Blue horizontal lines denote decrease of PCC/precuneus response compared to the baseline during the encoding periods and increase of PCC/precuneus response compared to the baseline during the recall periods (ps < 0.05).

Recall In contrast, suppression of high-gamma power in the PCC/precuneus was absent during the recall periods. Rather, high-gamma power in the PCC/precuneus was enhanced compared to the AI mostly during the 1390 – 1530 msec window prior to recall (ps < 0.05, Figure 3a).

AI response compared to PCC/precuneus during encoding and recall in the CATVFR task

We next examined high-gamma power in the CATVFR task. In this task, participants were presented with a list of words with consecutive pairs of words from a specific category (for example, JEANS-COAT, GRAPE-PEACH, etc.) and subsequently asked to recall as many as possible from the original list (Methods, Tables S1, S2b, S3b, Figure 1b) (Qasim, Mohan, Stein, & Jacobs, 2023).

Encoding High-gamma power in PCC/precuneus was suppressed compared to the AI during the 570 – 790 msec interval (ps < 0.05, Figure 3b).

Recall High-gamma power mostly did not differ between AI and PCC/precuneus prior to recall (ps > 0.05, Figure 3b).

AI response compared to PCC/precuneus during encoding and recall in the PALVCR task

The PALVCR task also consisted of three periods: encoding, delay, and recall (Methods, Tables S1, S2c, S3c, Figure 1c). During encoding, a list of word-pairs was visually presented, and then participants were asked to verbally recall the cued word from memory during the recall periods.

Encoding High-gamma power in PCC/precuneus was suppressed compared to the AI during the memory encoding period, during the 470 – 950 msec and 2010 – 2790 msec windows (ps < 0.05, Figure 3c).

Recall High-gamma power mostly did not differ between AI and PCC/precuneus prior to recall (ps > 0.05, Figure 3c).

AI response compared to PCC/precuneus during encoding and recall in the WMSM task

We next examined high-gamma power in the WMSM task. Participants performed multiple trials of a spatial memory task in a virtual navigation paradigm (Goyal et al., 2018; Jacobs et al., 2016; Lee et al., 2018) similar to the Morris water maze (Morris, 1984) (Methods, Tables S1, S2d, S3d, Figure 1d). Participants were shown objects in various locations during the encoding periods and asked to retrieve the location of the objects during the recall period.

Encoding High-gamma power in PCC/precuneus was suppressed compared to the AI, mostly during the 1390 – 2030 msec and 3150 – 4690 msec window (ps < 0.05, Figure 3d).

Recall High-gamma power mostly did not differ between AI and PCC/precuneus (ps > 0.05, Figure 3d).

Replication of increased high-gamma power in AI compared to PCC/precuneus across four memory tasks

Earlier version: We next used replication BF analysis to estimate the degree of replicability of high-gamma power suppression of the PCC/precuneus compared to the AI during the memory encoding periods of the four tasks. We used the posterior distribution obtained from the VFR (primary) dataset as a prior distribution for the test of data from the CATVFR, PALVCR, and WMSM (replication) datasets (Ly et al., 2019) (see Methods for details). We used the encoding time-windows for which we most consistently observed decrease of PCC/precuneus high-gamma power compared to the AI. These correspond to 110 – 1600 msec during the VFR task, 570 – 790 msec in the CATVFR task, 2010 – 2790 msec in the PALVCR task, and 3150 – 4690 msec in the WMSM task. We first averaged the high-gamma power across these strongest time-windows for each task and then used replication BF analysis to estimate the degree of replicability of high-gamma power suppression of the PCC/precuneus compared to the AI. Findings corresponding to the high-gamma power suppression of the PCC/precuneus compared to AI were replicated in the PALVCR (BF 5.16e+1) and WMSM (BF 2.69e+8) tasks. These results demonstrate very high replicability of high-gamma power suppression of the PCC/precuneus compared to AI during memory encoding. The consistent suppression effect was localized only to the PCC/precuneus, but not to the mPFC node of the DMN or the dPPC and MFG nodes of the FPN (Figures S1-S3).

Revised version: We next used replication BF analysis to estimate the degree of replicability of high-gamma power suppression of the PCC/precuneus compared to the AI during the memory encoding periods of the four tasks. We used the posterior distribution obtained from the VFR (primary) dataset as a prior distribution for the test of data from the CATVFR, PALVCR, and WMSM (replication) datasets (Ly et al., 2019) (see Methods for details). We used the encoding time-windows for which we most consistently observed decrease of PCC/precuneus high-gamma power compared to the AI. These correspond to 110 – 1600 msec during the VFR task, 570 – 790 msec in the CATVFR task, 2010 – 2790 msec in the PALVCR task, and 3150 – 4690 msec in the WMSM task. We first averaged the high-gamma power across these strongest time-windows for each task and then used replication BF analysis to estimate the degree of replicability of high-gamma power suppression of the PCC/precuneus compared to the AI.

Added in the revised version: Findings corresponding to the high-gamma power suppression of the PCC/precuneus compared to AI were replicated in the PALVCR (BF 5.16e+1) and WMSM (BF 2.69e+8) tasks. These results demonstrate very high replicability of high-gamma power suppression of the PCC/precuneus compared to AI during memory encoding. The consistent suppression effect was localized only to the PCC/precuneus, but not to the mPFC node of the DMN or the dPPC and MFG nodes of the FPN (Figures S1-S3).

In contrast to memory encoding, a similar analysis of high-gamma power did not reveal a consistent pattern of increased high-gamma power in AI and suppression of the PCC/precuneus across the four tasks during memory recall (Figure 3).

AI and PCC/precuneus response during encoding and recall compared to resting baseline

We examined whether AI and PCC/precuneus high-gamma power response during the encoding and recall periods are enhanced or suppressed when compared to the baseline periods. High-gamma power in the AI was increased compared to the resting baseline during both the encoding and recall periods, and across all four tasks (ps < 0.05, Figure 3). This suggests an enhanced role for the AI during both memory encoding and recall compared to resting baseline.

In contrast, high-gamma power in the PCC/precuneus was reduced compared to the resting baseline in three tasks – VFR, PALVCR, and WMSM – providing direct evidence for PCC/precuneus suppression during memory encoding (Figure 3). We did not find any increased high-gamma power activity in the PCC/precuneus, compared to the baseline, during memory retrieval (Figure 3). These results provide evidence for PCC/precuneus suppression compared to both the AI and resting baseline, during externally triggered stimuli during encoding.

High-gamma power for other brain areas compared to resting baseline were not consistent across tasks (Figures S1-S3).

Directed information flow from the AI to the DMN during encoding

Earlier version: We next examined directed information flow from the AI to the PCC/precuneus and mPFC nodes of the DMN, during the memory encoding periods of the VFR task. We used phase transfer entropy (PTE) (Lobier et al., 2014) to evaluate directed influences from the AI to the PCC/precuneus and mPFC and vice-versa. Informed by recent electrophysiology studies in nonhuman primates, which suggest that broadband field potentials activity, rather than narrowband, governs information flow in the brain (Davis, Muller, Martinez-Trujillo, Sejnowski, & Reynolds, 2020; Davis, Muller, & Reynolds, 2022), we examined PTE in a 0.5 to 80 Hz frequency spectrum to assess dynamic causal influences of the AI on the DMN.

Revised version: We next examined directed information flow from the AI to the PCC/precuneus and mPFC nodes of the DMN, during the memory encoding periods of the VFR task. We used phase transfer entropy (PTE) (Lobier et al., 2014) to evaluate directed influences from the AI to the PCC/precuneus and mPFC and vice-versa. Informed by recent electrophysiology studies in nonhuman primates, which suggest that broadband field potentials activity, rather than narrowband, governs information flow in the brain (Davis, Muller, Martinez-Trujillo, Sejnowski, & Reynolds, 2020; Davis, Muller, & Reynolds, 2022), we examined PTE in a 0.5 to 80 Hz frequency spectrum to assess dynamic directed influences of the AI on the DMN.

Directed information flow from the AI to the PCC/precuneus (F(1, 264) = 59.36, p<0.001, Cohen’s d = 0.95) and mPFC (F(1, 208) = 13.96, p<0.001, Cohen’s d = 0.52) were higher, than the reverse (Figure 4a).

Earlier version: Causal directed information flow between the anterior insula and the PCC/precuneus and mPFC nodes of the default mode network (DMN), across verbal and spatial memory domains, measured using phase transfer entropy (PTE).

Revised version: Directed information flow between the anterior insula and the PCC/precuneus and mPFC nodes of the default mode network (DMN), across verbal and spatial memory domains, measured using phase transfer entropy (PTE). (a) Experiment 1, VFR:

Earlier version: (a) Experiment 1, VFR: The anterior insula showed higher causal directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=142) during both encoding and recall. The anterior insula also showed higher causal directed information flow to the mPFC (AI → mPFC) compared to the reverse direction (mPFC → AI) (n=112) during both memory encoding and recall. (b) Experiment 2, CATVFR: The anterior insula showed higher causal directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=46) during both encoding and recall. (c) Experiment 3, PALVCR: The anterior insula showed higher causal directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=10) during both encoding and recall. (d) Experiment 4, WMSM: The anterior insula showed higher causal directed information flow to PCC/precuneus (AI → PCC/Pr) than the reverse (PCC/Pr → AI) (n=91), during both spatial memory encoding and recall. The anterior insula also showed higher causal directed information flow to mPFC (AI → mPFC) than the reverse (mPFC → AI) (n=23), during both spatial memory encoding and recall. In each panel, the direction for which PTE is higher, is underlined. White dot in each violin plot represents median PTE across electrode pairs. *** p < 0.001, * p < 0.05.

Revised version: The anterior insula showed higher directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=142) during both encoding and recall. The anterior insula also showed higher directed information flow to the mPFC (AI → mPFC) compared to the reverse direction (mPFC → AI) (n=112) during both memory encoding and recall. (b) Experiment 2, CATVFR: The anterior insula showed higher directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=46) during both encoding and recall. (c) Experiment 3, PALVCR: The anterior insula showed higher directed information flow to the PCC/precuneus (AI → PCC/Pr) compared to the reverse direction (PCC/Pr → AI) (n=10) during both encoding and recall. (d) Experiment 4, WMSM: The anterior insula showed higher directed information flow to PCC/precuneus (AI → PCC/Pr) than the reverse (PCC/Pr → AI) (n=91), during both spatial memory encoding and recall. The anterior insula also showed higher directed information flow to mPFC (AI → mPFC) than the reverse (mPFC → AI) (n=23), during both spatial memory encoding and recall. In each panel, the direction for which PTE is higher, is underlined. White dot in each violin plot represents median PTE across electrode pairs. *** p < 0.001, * p < 0.05.

Earlier version: Replication across three experiments with BF We used replication BF analysis to estimate the degree of replicability of direction of information flow across the four experiments (Table 1a, Figures 4b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments). Findings corresponding to the direction of information flow between the AI and the PCC/precuneus during memory encoding were replicated all three tasks (BFs 9.31e+5, 1.44e+4, and 1.68e+18 for CATVFR, PALVCR, and WMSM respectively). Findings corresponding to the direction of information flow between the AI and mPFC during memory encoding were also replicated across all three tasks (BFs 4.10e+1, 8.78e+0, and 5.34e+5 for CATVFR, PALVCR, and WMSM respectively). This highly consistent pattern of results was not observed in any other frequency band (delta-theta (0.5-8 Hz), alpha (8-12 Hz), beta (12-30 Hz), gamma (30-80 Hz), and high-gamma (80-160 Hz); results not shown). These results demonstrate very high replicability of directed information flow from the AI to the DMN nodes during memory encoding.

Revised version: Replication across three experiments with BF We used replication BF analysis to estimate the degree of replicability of direction of information flow across the four experiments (Table 1a, Figures 4b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments). Findings corresponding to the direction of information flow between the AI and the PCC/precuneus during memory encoding were replicated all three tasks (BFs 9.31e+5, 1.44e+4, and 1.68e+18 for CATVFR, PALVCR, and WMSM respectively).

Added in the revised version: Findings corresponding to the direction of information flow between the AI and mPFC during memory encoding were also replicated across all three tasks (BFs 4.10e+1, 8.78e+0, and 5.34e+5 for CATVFR, PALVCR, and WMSM respectively). This highly consistent pattern of results was not observed in any other frequency band (delta-theta (0.5-8 Hz), alpha (8-12 Hz), beta (12-30 Hz), gamma (30-80 Hz), and high-gamma (80-160 Hz); results not shown). These results demonstrate very high replicability of directed information flow from the AI to the DMN nodes during memory encoding.

These results demonstrate robust directed information flow from the AI to the PCC/precuneus and mPFC nodes of the DMN during memory encoding.

Directed information flow from the AI to the DMN during recall

Next, we examined directed influences of the AI on PCC/precuneus and mPFC during the recall phase of the verbal episodic memory task. During memory recall, directed information flow from the AI to the PCC/precuneus (F(1, 264) = 43.09, p<0.001, Cohen’s d = 0.81) and mPFC (F(1, 211) = 21.94, p<0.001, Cohen’s d = 0.65) were higher, than the reverse (Figure 4a).

Replication across three experiments with BF We next repeated the replication BF analysis for the recall periods of the memory tasks (Table 1b, Figures 4b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments). Findings corresponding to the direction of information flow between the AI and the PCC/precuneus during memory recall were replicated across all three tasks (BFs 1.30e+5, 6.74e+0, and 2.54e+10 for CATVFR, PALVCR, and WMSM respectively). Findings corresponding to the direction of information flow between the AI and the mPFC during memory recall were also replicated across the CATVFR and WMSM tasks (BFs 2.02e+1 and 1.32e+4 respectively).

These results demonstrate very high replicability of directed information flow from the AI to the DMN nodes across verbal and spatial memory tasks, during both memory encoding and recall.

Directed information flow from AI to FPN nodes during memory encoding

We next probed directed information flow between the AI and FPN nodes during the encoding periods of the verbal free recall task. Directed information flow from the AI to the dPPC (F(1, 1143) = 11.69, p<0.001, Cohen’s d = 0.20) and MFG (F(1, 1245) = 21.69, p<0.001, Cohen’s d = 0.26) were higher, than the reverse during memory encoding of the VFR task (Figure 5a).

Earlier version: Causal directed information flow between the anterior insula and the dPPC and MFG nodes of the frontoparietal network (FPN), across verbal and spatial memory domains.

Revised version: Directed information flow between the anterior insula and the dPPC and MFG nodes of the frontoparietal network (FPN), across verbal and spatial memory domains. (a) Experiment 1, VFR:

Earlier version: (a) Experiment 1, VFR: The anterior insula showed higher causal directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=586) during both encoding and recall. The anterior insula also showed higher causal directed information flow to the MFG (AI → MFG) compared to the reverse direction (MFG → AI) (n=642) during both memory encoding and recall. (b) Experiment 2, CATVFR: The anterior insula showed higher causal directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=327) during both encoding and recall. (c) Experiment 3, PALVCR: The anterior insula showed higher causal directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=242) during both encoding and recall. The anterior insula also showed higher causal directed information flow to the MFG (AI → MFG) compared to the reverse direction (MFG → AI) (n=362) during memory recall. (d) Experiment 4, WMSM: The anterior insula showed higher causal directed information flow to MFG (AI → MFG) than the reverse (MFG → AI) (n=177), during both spatial memory encoding and recall. In each panel, the direction for which PTE is higher, is underlined. *** p < 0.001, ** p < 0.01.

Revised version: The anterior insula showed higher directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=586) during both encoding and recall. The anterior insula also showed higher directed information flow to the MFG (AI → MFG) compared to the reverse direction (MFG → AI) (n=642) during both memory encoding and recall. (b) Experiment 2, CATVFR: The anterior insula showed higher directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=327) during both encoding and recall. (c) Experiment 3, PALVCR: The anterior insula showed higher directed information flow to the dorsal PPC (AI → dPPC) compared to the reverse direction (dPPC → AI) (n=242) during both encoding and recall. The anterior insula also showed higher directed information flow to the MFG (AI → MFG) compared to the reverse direction (MFG → AI) (n=362) during memory recall. (d) Experiment 4, WMSM: The anterior insula showed higher directed information flow to MFG (AI → MFG) than the reverse (MFG → AI) (n=177), during both spatial memory encoding and recall. In each panel, the direction for which PTE is higher, is underlined. *** p < 0.001, ** p < 0.01.

Replication across three experiments with BF We used replication BF analysis for the replication of AI directed influences on FPN nodes during the encoding phase of the memory tasks (Table 1a, Figures 5b-d, Supplementary Results). Similarly, we also obtained very high BFs for findings corresponding to the direction of information flow between the AI and dPPC (BFs > 2.33e+26) and also between the AI and MFG (BFs > 2.35e+27), across all three tasks.

These results demonstrate that the AI has robust directed information flow to the dPPC and MFG nodes of the FPN during memory encoding.

Directed information flow from AI to FPN nodes during memory recall

Directed influences from the AI to the dPPC (F(1, 1143) = 17.47, p<0.001, Cohen’s d = 0.25) and MFG (F(1, 1246) = 42.75, p<0.001, Cohen’s d = 0.37) were higher, than the reverse during memory recall of the VFR task (Figure 5a).

Replication across three experiments with BF We also found very high BFs for findings corresponding to the direction of information flow between the AI and the dPPC (BFs > 4.51e+27) and MFG (BFs > 6.90e+27) nodes of the FPN across the CATVFR, PALVCR, and WMSM tasks, during the memory recall period (Table 1b, Figures 5b-d, Supplementary Results).

These results demonstrate very high replicability of directed information flow from the AI to the FPN nodes across multiple memory experiments, during both memory encoding and recall.

Comparison of directed information flow: AI vs. Inferior Frontal Gyrus

To examine the specificity of the AI directed information flow to the DMN and FPN, we conducted a control analysis using electrodes implanted in the inferior frontal gyrus (IFG, BA 44). The IFG serves as an ideal control region due to its anatomical adjacency to the AI, its involvement in a wide range of cognitive control functions including response inhibition (Cai, Ryali, Chen, Li, & Menon, 2014), and its frequent co-activation with the AI in fMRI studies. Furthermore, the IFG has been associated with controlled retrieval of memory (Badre et al., 2005; Badre & Wagner, 2007; Wagner et al., 2001), making it a compelling region for comparison.

Earlier version: Our analysis revealed a striking contrast between the AI and IFG in their patterns of directed information flow. While the AI exhibited strong causal influences on both the DMN and FPN, the IFG showed the opposite pattern. Specifically, both the DMN and FPN demonstrated higher influence on the IFG than the reverse, during both encoding and recall periods, and across all four memory experiments (Figures S4, S5).

Revised version: Our analysis revealed a striking contrast between the AI and IFG in their patterns of directed information flow. While the AI exhibited strong directed influences on both the DMN and FPN, the IFG showed the opposite pattern. Specifically, both the DMN and FPN demonstrated higher influence on the IFG than the reverse, during both encoding and recall periods, and across all four memory experiments (Figures S4, S5).

Added in the revised version: To quantify this difference more precisely, we calculated the net outflow for both regions, defined as the difference (PTE(out)–PTE(in), see Methods for details). This analysis revealed that the AI’s net outflow was significantly higher than that of the IFG during both encoding and recall phases, a finding replicated across all four experiments (all ps < 0.001) (Figure S6).

Added in the revised version: These results not only highlight the unique role of the AI in orchestrating large-scale network dynamics during memory processes but also demonstrate the specificity of this function when compared to an anatomically adjacent and functionally relevant region. The consistent pattern across diverse memory tasks and experimental phases underscores the robustness of the AI’s role as an outflow hub during memory formation and retrieval.

Removed in the revised version: These findings highlight the unique role of the AI in orchestrating large-scale network dynamics during memory processes. The AI’s pattern of directed information flow stands in contrast to that of the IFG, despite their anatomical proximity and shared involvement in cognitive control processes. This dissociation underscores the specificity of the AI’s function in coordinating network interactions during memory formation and retrieval.

Enhanced information flow from the AI to the DMN and FPN during episodic memory processing, compared to resting-state baseline

Earlier version: We next examined whether directed information flow from the AI to the DMN and FPN nodes during the memory tasks differed from the resting-state baseline. Resting-state baselines were extracted immediately before the start of the task sessions and the duration of task and rest epochs were matched to ensure that differences in network dynamics could not be explained by differences in duration of the epochs. Directed information flow from the AI to both the DMN and FPN were higher during both the memory encoding and recall phases and across the four experiments, compared to baseline in all but two cases (Figures S6, S7). These findings provide strong evidence for enhanced role of AI directed information flow to the DMN and FPN during memory processing compared to the resting state.

Revised version: We next examined whether directed information flow from the AI to the DMN and FPN nodes during the memory tasks differed from the resting-state baseline. Resting-state baselines were extracted immediately before the start of the task sessions and the duration of task and rest epochs were matched to ensure that differences in network dynamics could not be explained by differences in duration of the epochs. Directed information flow from the AI to both the DMN and FPN were higher during both the memory encoding and recall phases and across the four experiments, compared to baseline in all but two cases (Figures S7, S8).

Added in the revised version: To further elucidate the task-specific role of the anterior insula (AI), we compared its net outward directed influence during memory tasks to that observed during resting state. We quantified this influence as the difference between outgoing and incoming information flow (PTE(out) - PTE(in)). This analysis revealed that the AI’s net outflow was significantly enhanced during both encoding and recall phases of memory tasks compared to resting state in all but one case (ps < 0.05) (Figure S9). This pattern was consistently observed across all four experiments. These findings provide strong evidence for enhanced role of AI directed information flow to the DMN and FPN during memory processing compared to the resting state.

Differential information flow from the AI to the DMN and FPN for successfully recalled and forgotten memory trials

Earlier version: We examined memory effects by comparing PTE between successfully recalled and forgotten memory trials. However, this analysis did not reveal differences in directed influence from the AI on the DMN and FPN or the reverse, between successfully recalled and forgotten memory trials during the encoding as well as recall periods in any of the memory experiments (all ps>0.05).

Revised version: We examined memory effects by comparing PTE between successfully recalled and forgotten memory trials. However, this analysis did not reveal differences in directed influence from the AI on the DMN and FPN or the reverse, between successfully recalled and forgotten memory trials during the encoding as well as recall periods in any of the memory experiments (all ps>0.05) (Figures S10, S11).

Earlier version: Causal outflow hub during encoding and recall

Revised version: Outflow hub during encoding and recall

Earlier version: fMRI studies have suggested that the AI acts as an outflow hub with respect to interactions with the DMN and FPN (Sridharan et al., 2008). To test the potential neural basis of this finding, we calculated net outflow as the difference between the total outgoing information and total incoming information (PTE(out)–PTE(in), see Methods for details).

Revised version: fMRI studies have suggested that the AI acts as an outflow hub with respect to interactions with the DMN and FPN (Sridharan et al., 2008). To test the potential neural basis of this finding, we calculated net outflow (PTE(out)–PTE(in)) as the difference between the total outgoing information and total incoming information.

Encoding This analysis revealed that the net outflow from the AI is positive and higher than the PCC/precuneus (F(1, 3319) = 154.8, p<0.001, Cohen’s d = 0.43) node of the DMN, in the VFR task (Figure 6a).

Earlier version: The anterior insula is a causal outflow hub in its interactions with the DMN and FPN, during encoding and recall periods, and across memory experiments.

Revised version: The anterior insula is an outflow hub in its interactions with the DMN and FPN, during encoding and recall periods, and across memory experiments.

In each panel, the net direction of information flow between the AI and the DMN and FPN nodes are indicated by green arrows on the right. *** p < 0.001, ** p < 0.01, * p < 0.05.

This analysis also revealed that the net outflow from the AI is higher than both the dPPC (F(1, 5346) = 67.87, p<0.001, Cohen’s d = 0.23) and MFG (F(1, 6920) = 132.74, p<0.001, Cohen’s d = 0.28) nodes of the FPN, in the VFR task (Figure 6a).

Findings in the VFR task were also replicated across the CATVFR, PALVCR, and WMSM tasks, where we found that the net outflow from the AI is higher than the PCC/precuneus and mPFC nodes of the DMN and the dPPC and MFG nodes of the FPN (Figures 6b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments).

Recall Net outflow from the AI is positive and higher than both PCC/precuneus (F(1, 3287) = 151.21, p<0.001, Cohen’s d = 0.43) and mPFC (F(1, 4694) = 7.81, p<0.01, Cohen’s d = 0.08) during the recall phase of the VFR task (Figure 6a).

Net outflow from the AI is also higher than both the dPPC (F(1, 5388) = 90.71, p<0.001, Cohen’s d = 0.26) and MFG (F(1, 6945) = 167.14, p<0.001, Cohen’s d = 0.31) nodes of the FPN during recall (Figure 6a).

Earlier version: Crucially, these findings were also replicated across the CATVFR, PALVCR, and WMSM tasks and during both encoding and recall periods (Figures 6b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments). Together, these results demonstrate that the AI is a causal outflow hub in its interactions with the PCC/precuneus and mPFC nodes of the DMN and also the dPPC and MFG nodes of the FPN, during both verbal and spatial memory encoding and recall.

Revised version: Crucially, these findings were also replicated across the CATVFR, PALVCR, and WMSM tasks and during both encoding and recall periods (Figures 6b-d, also see Supplementary Results for detailed stats related to the CATVFR, PALVCR, and WMSM experiments). Together, these results demonstrate that the AI is an outflow hub in its interactions with the PCC/precuneus and mPFC nodes of the DMN and also the dPPC and MFG nodes of the FPN, during both verbal and spatial memory encoding and recall.

Narrowband phase synchronization between the AI and the DMN and FPN during encoding and recall compared to resting baseline

We next directly compared the phase locking values (PLVs) (see Methods for details) between the AI and the PCC/precuneus and mPFC nodes of the DMN and also the dPPC and MFG nodes of the FPN for the encoding and the recall periods compared to resting baseline. However, narrowband PLV values did not significantly differ between the encoding/recall vs. rest periods, in any of the delta-theta (0.5-8 Hz), alpha (8-12 Hz), beta (12-30 Hz), gamma (30-80 Hz), and high-gamma (80-160 Hz) frequency bands. These results indicate that PTE, rather than phase synchronization, more robustly captures the AI dynamic interactions with the DMN and the FPN.

Discussion

Our study investigated the electrophysiological underpinnings of large-scale brain network interactions during episodic memory processes, focusing on the dynamic interplay between the salience network (SN), default mode network (DMN), and frontoparietal network (FPN) as conceptualized in the triple network model (Cai et al., 2021; Menon, 2011, 2023). This model has been primarily investigated in the context of cognitive control tasks (Cai et al., 2021; Menon, 2011, 2023). However, its applicability to memory processes remains less explored, particularly at the electrophysiological level. We elucidated how these three networks interact during different phases of memory processing, focusing on the directed information flow between key cortical nodes. The triple network model posits distinct roles for each network: the SN, anchored by the anterior insula, is thought to detect behaviorally relevant stimuli and orient attention towards information that needs to be encoded; the DMN is implicated in internally-driven processes, and memory recall; and the FPN contributes to the maintenance and manipulation of information in working memory, processes critical for both encoding and recall (Badre et al., 2005; Badre & Wagner, 2007; Wagner et al., 2001; Wagner et al., 2005). By leveraging intracranial EEG data from a large cohort of participants across four diverse memory tasks, we sought to provide a comprehensive, high-temporal resolution account of these network dynamics.

Earlier version: We discovered that the anterior insula (AI), a crucial node of the SN, exerts strong directed influence on both the DMN and FPN during both memory encoding and recall. This finding was consistently observed across multiple experiments spanning verbal and spatial memory domains, highlighting the robustness and generalizability of our results. Importantly, our study extends the applicability of the triple network model beyond cognitive control tasks to episodic memory processes, thus broadening its explanatory power in the context of memory formation. Furthermore, we observed a distinctive suppression of high-gamma power in the posterior cingulate cortex/precuneus (PCC/precuneus) node of the DMN compared to the AI during memory encoding, suggesting a task-specific functional down-regulation of this region. Our findings significantly advance the understanding of the SN’s role in modulating large-scale brain networks during episodic memory formation and underscore the importance of the triple network model in domain-general coordination of brain networks (Figure 7).

Revised version: We discovered that the anterior insula (AI), a crucial node of the SN, exerts strong directed influence on both the DMN and FPN during both memory encoding and recall. This finding was consistently observed across multiple experiments spanning verbal and spatial memory domains, highlighting the robustness and generalizability of our results. Importantly, our study extends the applicability of the triple network model beyond cognitive control tasks to episodic memory processes, thus broadening its explanatory power in the context of memory formation.

Added in the revised version: Furthermore, we observed a distinctive suppression of high-gamma power in the posterior cingulate cortex/precuneus (PCC/precuneus) node of the DMN compared to the AI during memory encoding, suggesting a task-specific functional down-regulation of this region. Our findings significantly advance the understanding of the SN’s role in modulating large-scale brain networks during episodic memory formation and underscore the importance of the triple network model in domain-general coordination of brain networks (Figure 7).

Earlier version: Schematic illustration of key findings related to the intracranial electrophysiology of the triple network model in human episodic memory.

Revised version: Schematic illustration of key findings related to the intracranial electrophysiology of the triple network model in human episodic memory. (a) High-gamma response:

Earlier version: (a) High-gamma response: Our analysis of local neuronal activity revealed consistent suppression of high-gamma power in the PCC/precuneus compared to the AI during encoding periods across all four episodic memory experiments. We did not consistently observe any significant differences in high-gamma band power between AI and the mPFC node of the DMN or the dPPC and MFG nodes of the FPN during the encoding periods across the four episodic memory experiments. In contrast, we detected similar high-gamma band power in the PCC/precuneus relative to the AI during the recall periods. (b) Directed information flow: Despite variable patterns of local activation and suppression across DMN and FPN nodes during memory encoding and recall, we found stronger causal influence (denoted by green arrows, thickness of arrows denotes degree of replicability across experiments, see Table 1) by the AI on both the DMN as well as the FPN nodes than the reverse, across all four memory experiments, and during both encoding and recall periods.

Revised version: Our analysis of local neuronal activity revealed consistent suppression of high-gamma power in the PCC/precuneus compared to the AI during encoding periods across all four episodic memory experiments. We did not consistently observe any significant differences in high-gamma band power between AI and the mPFC node of the DMN or the dPPC and MFG nodes of the FPN during the encoding periods across the four episodic memory experiments. In contrast, we detected similar high-gamma band power in the PCC/precuneus relative to the AI during the recall periods. (b) Directed information flow: Despite variable patterns of local activation and suppression across DMN and FPN nodes during memory encoding and recall, we found stronger directed influence (denoted by green arrows, thickness of arrows denotes degree of replicability across experiments, see Table 1) by the AI on both the DMN as well as the FPN nodes than the reverse, across all four memory experiments, and during both encoding and recall periods.

Earlier version: Investigating causal inter-network interactions using iEEG and phase transfer entropy

Revised version: Investigating directed inter-network interactions using iEEG and phase transfer entropy

Dynamic interactions between the AI and the DMN and FPN are hypothesized to shape human cognition (Cai et al., 2016; Cai et al., 2014; Dosenbach, Fair, Cohen, Schlaggar, & Petersen, 2008; Dosenbach et al., 2006; Menon, 2015b; Menon & Uddin, 2010). Although fMRI research has suggested that the AI plays a pivotal role in the task-dependent engagement and disengagement of the DMN and FPN across diverse cognitive tasks (Menon & Uddin, 2010; Sridharan et al., 2008), the neuronal basis of these results or the possibility of their being artifacts arising from slow dynamics and regional variation in the hemodynamic response inherent to fMRI signals remained unclear. To address these ambiguities, our analysis focused on casual interactions involving the AI and leveraged the high temporal resolution of iEEG signals. By investigating the directionality of information flow, we aimed to overcome the temporal resolution limitations of fMRI signals, providing a more mechanistic understanding of the AI’s role in modulating the DMN and FPN during memory formation. To assess reproducibility, we scrutinized network interactions across four different episodic memory tasks involving verbal free recall, categorized verbal free recall, paired associates learning verbal cued recall, and water maze spatial episodic memory tasks (Solomon et al., 2019).

Earlier version: We employed Phase Transfer Entropy (PTE), a robust metric of nonlinear and nonstationary dynamics to investigate dynamic interactions between the AI and four key cortical nodes of the DMN and FPN. PTE assesses the ability of one time-series to predict future values of another, estimating time-delayed causal influences, and is superior to methods like phase locking or coherence as it captures nonlinear and nonstationary interactions (Bassett & Sporns, 2017; Hillebrand et al., 2016; Lobier et al., 2014). PTE offers a robust and powerful tool for characterizing information flow between brain regions based on phase coupling (Hillebrand et al., 2016; Lobier et al., 2014; Wang et al., 2017) and has been successfully utilized in our previous studies (Das, de Los Angeles, & Menon, 2022; Das & Menon, 2020, 2021, 2022b, 2023).

Revised version: We employed Phase Transfer Entropy (PTE), a robust metric of nonlinear and nonstationary dynamics to investigate dynamic interactions between the AI and four key cortical nodes of the DMN and FPN. PTE assesses the ability of one time-series to predict future values of another, estimating time-delayed directed influences, and is superior to methods like phase locking or coherence as it captures nonlinear and nonstationary interactions (Bassett & Sporns, 2017; Hillebrand et al., 2016; Lobier et al., 2014). PTE offers a robust and powerful tool for characterizing information flow between brain regions based on phase coupling (Hillebrand et al., 2016; Lobier et al., 2014; Wang et al., 2017) and has been successfully utilized in our previous studies (Das, de Los Angeles, & Menon, 2022; Das & Menon, 2020, 2021, 2022b, 2023).

Broadband directed influences of the AI on DMN and FPN

Earlier version: Informed by recent electrophysiology studies in nonhuman primates, which suggest that broadband field potentials activity, rather than narrowband, governs information flow in the brain (Davis et al., 2020; Davis et al., 2022), we first examined PTE in a 0.5 to 80 Hz frequency spectrum to assess dynamic causal influences of the AI on the DMN and FPN. Our analysis revealed that AI exerts stronger influences on the PCC/precuneus and mPFC nodes of the DMN than the reverse. A similar pattern also emerged for FPN nodes, with the AI displaying stronger directed influences on the dPPC and MFG, than the reverse. Crucially, this asymmetric pattern of directed information flow was replicated across all four memory tasks. Moreover, this pattern also held during the encoding and recall of memory phases of all four tasks.

Revised version: Informed by recent electrophysiology studies in nonhuman primates, which suggest that broadband field potentials activity, rather than narrowband, governs information flow in the brain (Davis et al., 2020; Davis et al., 2022), we first examined PTE in a 0.5 to 80 Hz frequency spectrum to assess dynamic directed influences of the AI on the DMN and FPN. Our analysis revealed that AI exerts stronger influences on the PCC/precuneus and mPFC nodes of the DMN than the reverse. A similar pattern also emerged for FPN nodes, with the AI displaying stronger directed influences on the dPPC and MFG, than the reverse. Crucially, this asymmetric pattern of directed information flow was replicated across all four memory tasks. Moreover, this pattern also held during the encoding and recall of memory phases of all four tasks.

Replicability across memory tasks

Added in the revised version: Replication, a critical issue in all of systems neuroscience, is particularly challenging in the field of intracranial EEG studies, where data acquisition from patients is inherently difficult.

Earlier version: Replication, a critical issue in all of systems neuroscience, is particularly challenging in the field of intracranial EEG studies, where data acquisition from patients is inherently difficult. Compounding this issue is the virtual absence of data sharing and the substantial complexities involved in collecting electrophysiological data across distributed brain regions (Das & Menon, 2022b). Consequently, one of our study’s major objectives was to reproduce our findings across multiple experiments, bridging verbal and spatial memory domains and task phases. To quantify the degree of replicability of our findings across these domains, we employed replication Bayes Factor (BF) analysis (Ly et al., 2019; Verhagen & Wagenmakers, 2014). Our analysis revealed very high replication BFs related to replication of information flow from the AI to the DMN and FPN (Table 1). Specifically, the BFs associated with the replication of direction of information flow between the AI and the DMN and FPN were decisive (BFs > 100), demonstrating consistent results across various memory tasks and contexts.

Revised version: Compounding this issue is the virtual absence of data sharing and the substantial complexities involved in collecting electrophysiological data across distributed brain regions (Das & Menon, 2022b). Consequently, one of our study’s major objectives was to reproduce our findings across multiple experiments, bridging verbal and spatial memory domains and task phases. To quantify the degree of replicability of our findings across these domains, we employed replication Bayes Factor (BF) analysis (Ly et al., 2019; Verhagen & Wagenmakers, 2014). Our analysis revealed very high replication BFs related to replication of information flow from the AI to the DMN and FPN (Table 1). Specifically, the BFs associated with the replication of direction of information flow between the AI and the DMN and FPN were decisive (BFs > 100), demonstrating consistent results across various memory tasks and contexts.

Added in the revised version: Task-specific enhancement of AI’s directed influence: Contrasts with IFG and resting state

Added in the revised version: Our analysis revealed a striking contrast between the AI and inferior frontal gyrus (IFG) in their patterns of directed information flow. While the AI exhibited strong directed influences on both the DMN and FPN, the IFG demonstrated an inverse relationship. Specifically, both the DMN and FPN exerted higher influence on the IFG than vice versa, a pattern that held consistent across both encoding and recall periods, and throughout all four memory experiments (Figures S4, S5). Our analysis also revealed that the AI’s net outflow was significantly higher than that of the IFG during both encoding and recall phases, a finding replicated across all four experiments.

Added in the revised version: Furthermore, we compared the AI’s net outward directed influence during memory tasks to that observed during resting state. This analysis showed that the AI’s net outflow was significantly enhanced during both encoding and recall phases of memory tasks compared to resting state, consistently across all four experiments. This task-specific enhancement suggests that the AI’s role in coordinating large-scale network dynamics is specifically amplified during memory processes.

Added in the revised version: These results not only highlight the unique role of the AI in orchestrating large-scale network dynamics during memory processes but also demonstrate the specificity of this function when compared to an anatomically adjacent and functionally relevant region implicated in cognitive control (Badre et al., 2005; Badre & Wagner, 2007; Cai et al., 2014; Wagner et al., 2001). The consistent pattern across diverse memory tasks and experimental phases underscores the robustness of the AI’s role in memory-related network interactions.

High-gamma power suppression in the PCC/precuneus during encoding, but not recall

Our analysis of local neuronal activity revealed a consistent and specific pattern of high-gamma power suppression in the PCC/precuneus compared to the AI during memory encoding across all four episodic memory tasks. This finding aligns with the typical deactivation of default mode network (DMN) nodes during attention-demanding tasks (Wen et al., 2013), while also extending our understanding of the DMN’s role in episodic memory formation (Buckner et al., 2008; Menon, 2023).

Earlier version: Importantly, this suppression effect was confined to the PCC/precuneus within the DMN, with no parallel reductions observed in the mPFC. Moreover, suppression of the PCC/precuneus was stronger when compared to the dorsal posterior parietal cortex (dPPC) and middle frontal gyrus (MFG) nodes of the FPN (Figures S8, S9). Bayesian replication analysis substantiated the high degree of replicability of this PCC/precuneus suppression effect across tasks (Bayes Factors > 5.16e+1). These findings extend previous fMRI studies reporting DMN suppression during attention to external stimuli (Bressler & Menon, 2010; Raichle et al., 2001; Seeley et al., 2007) and complement optogenetic research in rodents’ brains demonstrating AI-induced suppression of DMN regions (Menon et al., 2023).

Revised version: Importantly, this suppression effect was confined to the PCC/precuneus within the DMN, with no parallel reductions observed in the mPFC. Moreover, suppression of the PCC/precuneus was stronger when compared to the dorsal posterior parietal cortex (dPPC) and middle frontal gyrus (MFG) nodes of the FPN (Figures S12, S13). Bayesian replication analysis substantiated the high degree of replicability of this PCC/precuneus suppression effect across tasks (Bayes Factors > 5.16e+1). These findings extend previous fMRI studies reporting DMN suppression during attention to external stimuli (Bressler & Menon, 2010; Raichle et al., 2001; Seeley et al., 2007) and complement optogenetic research in rodents’ brains demonstrating AI-induced suppression of DMN regions (Menon et al., 2023).

High-gamma activity (80-160 Hz) is a reliable indicator of localized, task-related neural processing, often associated with synchronized activity of local neural populations and elevated neuronal spiking (Canolty & Knight, 2010). High-gamma activity (typically ranging from 80-160 Hz) has been reliably implicated in various cognitive tasks across sensory modalities, including visual (Lachaux et al., 2005; Tallon-Baudry, Bertrand, Hénaff, Isnard, & Fischer, 2005), auditory (Crone, Boatman, Gordon, & Hao, 2001; Edwards, Soltani, Deouell, Berger, & Knight, 2005), and across cognitive domains, including working memory (Canolty et al., 2006; Mainy et al., 2007) and episodic memory (Daitch & Parvizi, 2018; Sederberg et al., 2007). The suppression we observed during encoding likely reflects functional down-regulation of the PCC/precuneus, potentially to minimize interference from internally-oriented processes during the encoding of external information.

In contrast, during memory recall, we observed different patterns of activity. In the three verbal tasks (VFR, CATVFR, and PALVCR), PCC/precuneus activity showed enhanced responses compared to the AI in the 1-1.6 second window prior to word production. However, it’s crucial to note that our analysis was time-locked to word production rather than the onset of internal retrieval processes. In the spatial memory task WMSM, the PCC/precuneus exhibited an earlier onset and enhanced activity compared to the AI. This task may provide a clearer window into recall processes: findings align with the view that DMN nodes may play a crucial role in triggering internal recall processes. However, the precise timing of internal retrieval initiation remains a challenge in verbal tasks, potentially limiting our ability to capture the full dynamics of regional activity, and its replicability, during early stages of recall.

The observed high-gamma suppression in the PCC/precuneus during encoding, but not recall, likely reflects the distinct cognitive demands of these memory phases. Encoding primarily involves externally-driven processes, requiring attention to and processing of incoming stimuli. In contrast, recall is predominantly internally-driven, relying on the retrieval and reconstruction of stored information. This dissociation in PCC/precuneus activity aligns with its known role in the DMN, which typically shows deactivation during externally-oriented tasks and activation during internally-directed cognition. This pattern of activity underscores the flexible and context-dependent functioning of brain regions within large-scale networks, adapting their engagement to support different aspects of memory processing.

Broadband vs. high-gamma directed influences

Earlier version: Notably, our findings reveal a robust and consistent directed influence exerted by the AI on all nodes of both the DMN and the FPN, extending across all four memory tasks and both memory encoding and recall phases. These causal influences were prominently manifested in broadband signals. Interestingly, such directed influences were not observed in the high-gamma frequency range (80-160 Hz). This absence aligns with current models positing that high-gamma activity is more likely to reflect localized processing, while lower-frequency bands are implicated in longer-range network communication and coordination (Bastos et al., 2015; Das et al., 2022; Das & Menon, 2020, 2021, 2022b, 2023; K. J. Miller et al., 2007). More generally, our findings emphasize that it is crucial to differentiate between high-gamma activity (f > 80 Hz) and sub-high-gamma (f < 80 Hz) fluctuations, as these signal types are indicative of different underlying physiological processes, each with distinct implications for understanding neural network dynamics.

Revised version: Notably, our findings reveal a robust and consistent directed influence exerted by the AI on all nodes of both the DMN and the FPN, extending across all four memory tasks and both memory encoding and recall phases. These directed influences were prominently manifested in broadband signals. Interestingly, such directed influences were not observed in the high-gamma frequency range (80-160 Hz). This absence aligns with current models positing that high-gamma activity is more likely to reflect localized processing, while lower-frequency bands are implicated in longer-range network communication and coordination (Bastos et al., 2015; Das et al., 2022; Das & Menon, 2020, 2021, 2022b, 2023; K. J. Miller et al., 2007). More generally, our findings emphasize that it is crucial to differentiate between high-gamma activity (f > 80 Hz) and sub-high-gamma (f < 80 Hz) fluctuations, as these signal types are indicative of different underlying physiological processes, each with distinct implications for understanding neural network dynamics.

Successful and unsuccessful memory effects engage similar AI-directed circuits

Our analysis revealed no significant differences in directed connectivity between successfully recalled and forgotten memory trials, suggesting that the reported effects may not be specific to successful memory formation and may be related to attentional or other general cognitive processing rather than memory processing per se. While our study provides valuable insights into the interactions between the AI and the DMN and FPN during cognitive tasks involving verbal and spatial information processing during memory tasks, it is crucial to acknowledge that these interactions may not be unique to memory processes. The AI’s directed influence on the DMN and FPN could reflect a more general role in coordinating attentional resources, which are essential for various cognitive functions, including memory formation (Menon & Uddin, 2010; Uddin, 2015). To disentangle the specific contributions of memory recall and attention, future studies should incorporate carefully designed control tasks that do not involve memory components. It is also important to note that successful memory recall likely involves the coordinated activity of multiple brain systems beyond the triple network model investigated here. For instance, the medial temporal lobe, including the hippocampus and adjacent cortical regions, plays a crucial role in episodic memory formation and retrieval (Burgess, Maguire, & O’Keefe, 2002; Moscovitch et al., 2016). Future studies will need to investigate a broader set of brain areas during successful and unsuccessful memory trials to gain a more comprehensive understanding of the neural circuits supporting distinctions between successfully recalled and forgotten memory trials.

Externally triggered vs. internally driven memory processes

Our results reveal a consistent pattern of directed information flow from the AI to both the DMN and FPN, persisting across externally triggered encoding and internally driven free recall. This pattern underscores the AI’s robust and versatile role in modulating large-scale brain networks across diverse task contexts, aligning with the triple network model’s conceptualization of the AI as a critical hub for attentional and cognitive control (Menon, 2011, 2023). However, the persistence of AI-driven information flow during internally triggered free recall was unexpected, given the view of the DMN’s dominance in internal cognition. This reproducible pattern, observed across both externally and internally driven tasks in all four experiments, reinforces the AI’s crucial role in orchestrating network dynamics over extended time periods.

Earlier version: We did not detect an opposing pattern of greater causal influences from the DMN during recall, as might be expected given the internally-driven nature of free recall. Several factors may contribute to this unexpected result. First, in the three verbal recall tasks, our PTE analysis was time-locked to word production onset, which may not capture the dynamics of network interactions during recall, particularly in the early retrieval initiation stage whose precise onset is unknown. This limitation is especially relevant for understanding the DMN’s role, which might be more prominent in the initiation of recall rather than the selection of verbal output. Secondly, the PTE method requires relatively long time series for robust estimation of information flow. The brief windows associated with initiation of individual recall events may not provide sufficient data for detecting subtle shifts in network dynamics, potentially masking transient increases in DMN influence.

Revised version: We did not detect an opposing pattern of greater directed influences from the DMN during recall, as might be expected given the internally-driven nature of free recall. Several factors may contribute to this unexpected result. First, in the three verbal recall tasks, our PTE analysis was time-locked to word production onset, which may not capture the dynamics of network interactions during recall, particularly in the early retrieval initiation stage whose precise onset is unknown. This limitation is especially relevant for understanding the DMN’s role, which might be more prominent in the initiation of recall rather than the selection of verbal output. Secondly, the PTE method requires relatively long time series for robust estimation of information flow.

Added in the revised version: The brief windows associated with initiation of individual recall events may not provide sufficient data for detecting subtle shifts in network dynamics, potentially masking transient increases in DMN influence.

Moreover, the consistent AI-driven information flow during recall might reflect the SN’s ongoing role in monitoring and evaluating retrieved information, even during internally-driven processes. This interpretation aligns with Sestieri and colleagues’ observation of sustained SN activity across all phases of memory search tasks (Sestieri et al., 2014) and suggests a more complex view of the AI’s function in both externally-driven and internal cognitive processes.

Earlier version: Intriguingly, as noted above, while we observed PCC/precuneus suppression during encoding and enhancement during recall, the AI maintained its causal influence on this DMN node during encoding and recall. This apparent discrepancy between local activity (suppression) and network-level communication highlights the complex nature of brain network dynamics. It is likely that PTE-based network interactions examined in this study at the time scale of about 2 seconds misses subtle changes in causal interactions that occur during internally-driven initiation of memory recall. Furthermore, our directed connectivity analysis used broadband signals (0.5-80 Hz), while power analysis of local neuronal activity focused on the high-gamma band (80-160 Hz). These different frequency ranges may capture distinct aspects of neural processing, with broadband connectivity reflecting more general, sustained network interactions.

Revised version: Intriguingly, as noted above, while we observed PCC/precuneus suppression during encoding and enhancement during recall, the AI maintained its directed influence on this DMN node during encoding and recall. This apparent discrepancy between local activity (suppression) and network-level communication highlights the complex nature of brain network dynamics. It is likely that PTE-based network interactions examined in this study at the time scale of about 2 seconds misses subtle changes in directed interactions that occur during internally-driven initiation of memory recall. Furthermore, our directed connectivity analysis used broadband signals (0.5-80 Hz), while power analysis of local neuronal activity focused on the high-gamma band (80-160 Hz). These different frequency ranges may capture distinct aspects of neural processing, with broadband connectivity reflecting more general, sustained network interactions.

To further elucidate these dynamics, future studies should consider employing techniques that can capture rapid changes in directed network interactions, investigating the temporal evolution of network interactions leading up to and following recall events, exploring the relationship between different frequency bands in connectivity and local activity measures, and developing methods to better estimate the onset of internal retrieval processes in verbal tasks. These approaches could provide valuable insights into the transition between externally-driven and internally-driven processes and offer a more precise understanding of the AI and PCC/precuneus’s differential roles in coordinating network dynamics across different memory phases.

AI as an outflow hub and a novel perspective on theoretical models of memory

Beyond information flow along individual pathways linking the AI with the DMN and FPN, our PTE analysis further revealed that the AI is an outflow hub in its interactions with the DMN and the FPN regardless of stimulus materials. As a central node of the salience network (Menon & Uddin, 2010; Seeley et al., 2007; Sridharan et al., 2008), the AI is known to play a crucial role in influencing other networks (Menon & Uddin, 2010; Uddin, 2015). Our results align with findings based on control theory analysis of brain networks during a working memory task. Specifically, Cai et al found higher causal outflow and controllability associated with the AI compared to DMN and FPN nodes during an n-back working memory task (Cai et al., 2021). Controllability refers to the ability to perturb a system from a given initial state to other configuration states in finite time by means of external control inputs. Intuitively, nodes with higher controllability require lower energy for perturbing a system from a given state, making controllability measures useful for identifying driver nodes with the potential to influence overall state dynamics. By virtue of its higher controllability relative to other brain areas, the AI is well-positioned to dynamically engage and disengage with other brain areas. These findings expand our understanding of the AI’s role, extending beyond attention and working memory tasks to incorporate two distinct stages of episodic memory formation. Our study, leveraging the temporal precision of iEEG data, substantially enhances previous fMRI findings by unveiling the neurophysiological mechanisms underlying the AI’s dynamic regulation of network activity during memory formation and cognition more generally.

Our findings bring a novel perspective to the seminal model of human memory proposed by Atkinson and Shiffrin (Atkinson & Shiffrin, 1968). This model conceptualizes memory as a multistage process, with control mechanisms regulating the transition of information across these stages. The observed suppression of high-gamma power in the PCC/precuneus and enhancement in the AI during the encoding phase may be seen as one neurophysiological manifestation of these control processes. The AI’s role as a dynamic switch, modulating activity between the DMN and FPN, aligns with active processing and control needed to encode sensory information into short-term memory. On the other hand, the transformations observed during the recall phase, particularly the discernable lack of DMN suppression patterns, may correspond to the retrieval processes where internally generated cues steer the reactivation of memory representations during recall. These results provide a novel neurophysiological model for understanding the complex control processes underpinning human memory functioning.

Limitations and future work

Earlier version: Our study, while revealing important insights into network dynamics during memory processes, has several limitations that provide avenues for further investigation. Although our computational methods suggest causal influences, direct causal manipulations, such as targeted brain stimulation during memory tasks, are needed to establish definitive causal relationships between network nodes. The PTE method, while powerful, cannot reliably capture rapid shifts in network dynamics. Subsequent research should employ techniques with higher temporal precision to map these changes.

Revised version: Our study, while revealing important insights into network dynamics during memory processes, has several limitations that provide avenues for further investigation. Although our computational methods suggest directed influences, direct causal manipulations, such as targeted brain stimulation during memory tasks, are needed to establish definitive causal relationships between network nodes. The PTE method, while powerful, cannot reliably capture rapid shifts in network dynamics. Subsequent research should employ techniques with higher temporal precision to map these changes.

To determine whether our observed network dynamics are memory-specific or reflect more general cognitive processes, additional work should compare directed connectivity patterns across memory and non-memory tasks. Our analysis approach, necessitated by limited multi-task participation, precluded robust within-subject analyses. Future studies should aim for more consistent multi-task participation to enable individual-level analyses of network dynamics across tasks.

In the free recall verbal tasks, precisely timing the onset of internal retrieval processes remains challenging. Experimental designs with cued recall similar to the WMSM task could provide crucial insights into early stages of memory retrieval. This approach could help clarify the roles of different networks, especially the DMN, during the initiation of recall versus the execution of verbal output. The dissociation we observed between local activity and network-level communication warrants further investigation. Further studies are needed to determine the relationship between different frequency bands in connectivity and local activity measures to better understand how these distinct aspects of neural processing contribute to memory formation and retrieval.

Despite these limitations, our findings provide a robust foundation for investigations into the electrophysiological basis of large-scale brain network interactions during memory formation and recall. By addressing these limitations, subsequent studies can further refine our understanding of how these networks dynamically coordinate to support episodic memory and other cognitive functions. Such investigations may reveal a more dynamic interplay between the SN, DMN, and FPN, where their relative influences shift rapidly depending on the specific cognitive demands of the task.

Conclusions

Earlier version: Our study provides novel insights into the neural dynamics underpinning episodic memory processes across four diverse memory experiments. We discovered that the anterior insula, a key node of the salience network, exerts a strong and consistent causal influence on both the default mode network and frontoparietal network during memory encoding and recall. This finding extends the applicability of the triple network model to episodic memory processes in both verbal and spatial domains, highlighting the anterior insula’s crucial role as a ‘causal hub’ that modulates information flow within and between these cognitive networks.

Revised version: Our study provides novel insights into the neural dynamics underpinning episodic memory processes across four diverse memory experiments. We discovered that the anterior insula, a key node of the salience network, exerts a strong and consistent directed influence on both the default mode network and frontoparietal network during memory encoding and recall. This finding extends the applicability of the triple network model to episodic memory processes in both verbal and spatial domains, highlighting the anterior insula’s crucial role as an outflow hub that modulates information flow within and between these cognitive networks.

Importantly, we observed a dissociation between local activity and network-level communication in the posterior cingulate cortex/precuneus node of the default mode network. The suppression of high-gamma power in this region during encoding, but not during recall, suggests a context-specific functional regulation that varies across memory phases. This finding reveals the intricate and dynamic interplay between local neural activity and large-scale network communication, and highlights the multifaceted nature of brain mechanisms underlying human memory processing.

The robust replicability of our findings across multiple memory tasks and modalities enhances the reliability and generalizability of our results, addressing a critical need in human intracranial EEG research. Our results reinforce the concept that memory operations rely on the concerted action of widely distributed brain networks (Mesulam, 1990), extending beyond traditional memory-specific regions.

By elucidating the electrophysiological basis of directed information flow within the triple network model, our study advances the understanding of neural circuit dynamics in human memory and cognition. Our findings provide a template for understanding the neural basis of memory impairments in neurological and psychiatric disorders. For instance, the disruption of these network interactions could contribute to memory deficits in conditions such as Alzheimer’s disease, where dysfunctions in the salience, default mode, and frontoparietal networks are now being increasingly documented (Bonthius, Solodkin, & Van Hoesen, 2005; Guzmán-Vélez et al., 2022).

Methods

UPENN-RAM iEEG recordings

iEEG recordings from 249 patients shared by Kahana and colleagues at the University of Pennsylvania (UPENN) (obtained from the UPENN-RAM public data release) were used for analysis (Jacobs et al., 2016). Patients with pharmaco-resistant epilepsy underwent surgery for removal of their seizure onset zones. iEEG recordings of these patients were downloaded from a UPENN-RAM consortium hosted data sharing archive (URL: http://memory.psych.upenn.edu/RAM). Prior to data collection, research protocols and ethical guidelines were approved by the Institutional Review Board at the participating hospitals and informed consent was obtained from the participants and guardians (Jacobs et al., 2016).

Details of all the recording sessions and data pre-processing procedures are described by Kahana and colleagues (Jacobs et al., 2016). Briefly, iEEG recordings were obtained using subdural grids and strips (contacts placed 10 mm apart) or depth electrodes (contacts spaced 5–10 mm apart) using recording systems at each clinical site. iEEG systems included DeltaMed XlTek (Natus), Grass Telefactor, and Nihon-Kohden EEG systems. Electrodes located in brain lesions or those which corresponded to seizure onset zones or had significant interictal spiking or had broken leads, were excluded from analysis.

Anatomical localization of electrode placement was accomplished by co-registering the postoperative computed CTs with the postoperative MRIs using FSL (FMRIB (Functional MRI of the Brain) Software Library), BET (Brain Extraction Tool), and FLIRT (FMRIB Linear Image Registration Tool) software packages. Preoperative MRIs were used when postoperative MRIs were not available. The resulting contact locations were mapped to MNI space using an indirect stereotactic technique and OsiriX Imaging Software DICOM viewer package.

We used the insula atlas by Faillenot and colleagues to demarcate the anterior insula (AI) (Faillenot et al., 2017), downloaded from http://brain-development.org/brain-atlases/adult-brain-atlases/. This atlas is based on probabilistic analysis of the anatomy of the insula with demarcations of the AI based on three short dorsal gyri and the posterior insula (PI) which encompasses two long gyri. To visualize iEEG electrodes on the insula atlas, we used surface-rendering code (GitHub: https://github.com/ludovicbellier/InsulaWM) provided by Llorens and colleagues (Llorens et al., 2023). We used the Brainnetome atlas (Fan et al., 2016) to demarcate the posterior cingulate cortex (PCC)/precuneus, the medial prefrontal cortex (mPFC), the dorsal posterior parietal cortex (dPPC), and the middle frontal gyrus (MFG). The dorsal anterior cingulate cortex node of the salience network was excluded from analysis due to lack of sufficient electrode placement. Out of 249 individuals, data from 177 individuals (aged from 16 to 64, mean age 36.3 ± 11.5, 91 females) were used for subsequent analysis based on electrode placement in the AI and the PCC/precuneus, mPFC, dPPC, and MFG.

Original sampling rates of iEEG signals were 500 Hz, 1000 Hz, 1024 Hz, and 1600 Hz. Hence, iEEG signals were downsampled to 500 Hz, if the original sampling rate was higher, for all subsequent analysis. The two major concerns when analyzing interactions between closely spaced intracranial electrodes are volume conduction and confounding interactions with the reference electrode (Burke et al., 2013; Frauscher et al., 2018). Hence bipolar referencing was used to eliminate confounding artifacts and improve the signal-to-noise ratio of the neural signals, consistent with previous studies using UPENN-RAM iEEG data (Burke et al., 2013; Ezzyat et al., 2018). Signals recorded at individual electrodes were converted to a bipolar montage by computing the difference in signal between adjacent electrode pairs on each strip, grid, and depth electrode and the resulting bipolar signals were treated as new “virtual” electrodes originating from the midpoint between each contact pair, identical to procedures in previous studies using UPENN-RAM data (Solomon et al., 2019). Line noise (60 Hz) and its harmonics were removed from the bipolar signals using band-stop filters at 57-63 Hz, 117-123 Hz, and 177-183 Hz. Finally, each bipolar signal was Z-normalized by removing mean and scaling by the standard deviation. For filtering, we used a fourth order two-way zero phase lag Butterworth filter throughout the analysis. iEEG signals were filtered in the broad frequency spectrum (0.5-80 Hz) as well as narrowband frequency spectra delta-theta (0.5-8 Hz), alpha (8-12 Hz), beta (12-30 Hz), gamma (30-80 Hz), and high-gamma (80-160 Hz).

Episodic memory experiments

Verbal free recall (VFR) task

Patients performed multiple trials of a verbal free recall experiment, where they were presented with a list of words and subsequently asked to recall as many as possible from the original list (Figure 1a) (Solomon et al., 2017; Solomon et al., 2019). The task consisted of three periods: encoding, delay, and recall. During encoding, a list of 12 words was visually presented for ∼30 sec. Words were selected at random, without replacement, from a pool of high frequency English nouns (http://memory.psych.upenn.edu/Word_Pools). Each word was presented for a duration of 1600 msec, followed by an inter-stimulus interval of 800 to 1200 msec. After the encoding period, participants engaged in a math distractor task (the delay period in Figure 1a), where they were instructed to solve a series of arithmetic problems in the form of a + b + c = ??, where a, b, and c were randomly selected integers ranging from 1 to 9. Mean accuracy across patients in the math task was 90.87% ± 7.22%, indicating that participants performed the math task with a high level of accuracy, similar to our previous studies (Das & Menon, 2022a). After a 20 sec post-encoding delay, participants were instructed to recall as many words as possible during the 30 sec recall period. Average recall accuracy across patients was 25.0% ± 10.6%, similar to prior studies of verbal episodic memory retrieval in neurosurgical patients (Burke et al., 2014). We analyzed iEEG epochs from the encoding and recall periods of the verbal free recall task. For the recall periods, iEEG recordings 1600 msec prior to the vocal onset of each word were analyzed (Solomon et al., 2019). Data from each trial was analyzed separately and specific measures were averaged across trials.

Categorized verbal free recall (CATVFR) task

This task was very similar to the verbal free recall task. Here, patients performed multiple trials of a categorized free recall experiment, where they were presented with a list of words with consecutive pairs of words from a specific category (for example, JEANS-COAT, GRAPE-PEACH, etc.) and subsequently asked to recall as many as possible from the original list (Figure 1b) (Qasim et al., 2023). Similar to the uncategorized verbal free recall task, this task also consisted of three periods: encoding, delay, and recall. During encoding, a list of 12 words was visually presented for ∼30 sec. Semantic categories were chosen using Amazon Mechanical Turk. Pairs of words from the same semantic category were never presented consecutively. Each word was presented for a duration of 1600 msec, followed by an inter-stimulus interval of 750 to 1000 msec. After a 20 sec post-encoding delay (math) similar to the uncategorized verbal free recall task, participants were instructed to recall as many words as possible during the 30 sec recall period. Average accuracy across patients in the math task was 89.46% ± 9.90%. Average recall accuracy across patients was 29.6% ± 13.4%. Analysis of iEEG epochs from the encoding and recall periods of the categorized free recall task was same as the uncategorized verbal free recall task.

Paired associates learning verbal cued recall (PALVCR) task

Patients performed multiple trials of a paired associates learning verbal cued recall experiment, where they were presented with a list of word-pairs and subsequently asked to recall based on the given word-cue (Figure 1c). Similar to the uncategorized verbal free recall task, this task also consisted of three periods: encoding, delay, and recall. During encoding, a list of 6 word-pairs was visually presented for ∼36 sec. Similar to the uncategorized verbal free recall task, words were selected at random, without replacement, from a pool of high frequency English nouns (http://memory.psych.upenn.edu/Word_Pools). Each word was presented for a duration of 4000 msec, followed by an inter-stimulus interval of 1750 to 2000 msec. After a 20 sec post-encoding delay (math) similar to the uncategorized verbal free recall task, participants were shown a specific word-cue for a duration of 4000 msec and asked to verbally recall the cued word from memory. Each word presentation during recall was followed by an inter-stimulus interval of 1750 to 2000 msec and the recall period lasted for ∼36 sec. Average accuracy across patients in the math task was 93.91% ± 4.66%. Average recall accuracy across patients was 33.8% ± 25.9%. For encoding, iEEG recordings corresponding to the 4000 msec encoding period of the task were analyzed. For recall, iEEG recordings 1600 msec prior to the vocal onset of each word were analyzed (Solomon et al., 2019). Data from each trial was analyzed separately and specific measures were averaged across trials.

Water maze spatial memory (WMSM) task

Patients performed multiple trials of a spatial memory experiment in a virtual navigation paradigm (Goyal et al., 2018; Jacobs et al., 2016; Lee et al., 2018) similar to the Morris water maze (Morris, 1984). The environment was rectangular (1.8:1 aspect ratio), and was surrounded by a continuous boundary (Figure 1d). There were four distal visual cues (landmarks), one centered on each side of the rectangle, to aid with orienting. Each trial (96 trials per session, 1–3 sessions per subject) started with two 5 sec encoding periods, during which subjects were driven to an object from a random starting location. At the beginning of an encoding period, the object appeared and, over the course of 5 sec, the subject was automatically driven directly toward it. The 5 sec period consisted of three intervals: first, the subject was rotated toward the object (1 sec), second, the subject was driven toward the object (3 sec), and, finally, the subject paused while at the object location (1 sec). After a 5 sec delay with a blank screen, the same process was repeated from a different starting location. After both encoding periods for each item, there was a 5 sec pause followed by the recall period. The subject was placed in the environment at a random starting location with the object hidden and then asked to freely navigate using a joystick to the location where they thought the object was located. When they reached their chosen location, they pressed a button to record their response. They then received feedback on their performance via an overhead view of the environment showing the actual and reported object locations. Average recall accuracy across patients was 48.1% ± 5.6%.

We analyzed the 5 sec iEEG epochs corresponding to the entire encoding and recall periods of the task as has been done previously (Goyal et al., 2018; Jacobs et al., 2016; Lee et al., 2018). Data from each trial was analyzed separately and specific measures were averaged across trials, similar to the verbal tasks.

Out of total 177 participants, 51% (91 out of 177) of participants participated in at least 2 experiments, 17% (30 out of 177) of participants participated in at least 3 experiments, and 6% (10 out of 177) of participants participated in all four experiments.

iEEG analysis of high-gamma power

We first filtered the signals in the high-gamma (80-160 Hz) frequency band (Canolty et al., 2006; Helfrich & Knight, 2016; Kai J. Miller, Weaver, & Ojemann, 2009) using sequential band-pass filters in increments of 10 Hz (i.e., 80–90 Hz, 90–100 Hz, etc.), using a fourth order two-way zero phase lag Butterworth filter. We used these narrowband filtering processing steps to correct for the 1/f decay of power. We then calculated the amplitude (envelope) of each narrow band signal by taking the absolute value of the analytic signal obtained from the Hilbert transform (Foster, Rangarajan, Shirer, & Parvizi, 2015). Each narrow band amplitude time series was then normalized to its own mean amplitude, expressed as a percentage of the mean. Finally, we calculated the mean of the normalized narrow band amplitude time series, producing a single amplitude time series. Signals were then smoothed using 0.2s windows with 90% overlap (Kwon et al., 2021) and normalized with respect to 0.2s pre-stimulus periods by subtracting the pre-stimulus baseline from the post-stimulus signal.

Earlier version: iEEG analysis of phase transfer entropy (PTE) and causal dynamics

Revised version: iEEG analysis of phase transfer entropy (PTE)

Earlier version: Phase transfer entropy (PTE) is a nonlinear measure of the directionality of information flow between time-series and can be applied as a measure of causality to nonstationary time-series (Das & Menon, 2021, 2022b; Lobier et al., 2014). Note that information flow described here relates to signaling between brain areas and does not necessarily reflect the representation or coding of behaviorally relevant variables per se. The PTE measure is in contrast to the Granger causality measure which can be applied only to stationary time-series (Barnett & Seth, 2014). We first carried out a stationarity test of the iEEG recordings (unit root test for stationarity (Barnett & Seth, 2014)) and found that the spectral radius of the autoregressive model is very close to one, indicating that the iEEG time-series is nonstationary. This precluded the applicability of the Granger causality analysis in our study.

Revised version: Phase transfer entropy (PTE) is a nonlinear measure of the directionality of information flow between time-series and can be applied to nonstationary time-series (Das & Menon, 2021, 2022b; Lobier et al., 2014). Note that information flow described here relates to signaling between brain areas and does not necessarily reflect the representation or coding of behaviorally relevant variables per se. The PTE measure is in contrast to the Granger causality measure which can be applied only to stationary time-series (Barnett & Seth, 2014). We first carried out a stationarity test of the iEEG recordings (unit root test for stationarity (Barnett & Seth, 2014)) and found that the spectral radius of the autoregressive model is very close to one, indicating that the iEEG time-series is nonstationary. This precluded the applicability of the Granger causality analysis in our study.

Earlier version: Given two time-series {xi } and {yi }, where i = 1, 2,…, M, instantaneous phases were first extracted using the Hilbert transform. Let and , where i = 1, 2,…, M, denote the corresponding phase time-series. If the uncertainty of the target signal at delay τ is quantified using Shannon entropy, then the PTE from driver signal to target signal can be given by where the probabilities can be calculated by building histograms of occurrences of singles, pairs, or triplets of instantaneous phase estimates from the phase time-series (Hillebrand et al., 2016). For our analysis, the number of bins in the histograms was set as 3.49 × STD × M −1/3 and delay τwas set as 2M / M±, where STD is average standard deviation of the phase time-series and and M ± is the number of times the phase changes sign across time and channels (Hillebrand et al., 2016). PTE has been shown to be robust against the choice of the delay τ and the number of bins for forming the histograms (Hillebrand et al., 2016). In our analysis, PTE was calculated for the entire encoding and recall periods for each trial and then averaged across trials.

Revised version: Given two time-series {xi } and {yi }, where i = 1, 2,…, M, instantaneous phases were first extracted using the Hilbert transform. Let , where i = 1, 2,…, M, denote the corresponding phase time-series. If the uncertainty of the target signal at delay 1″ is quantified using Shannon entropy, then the PTE from driver signal to target signal can be given by where the probabilities can be calculated by building histograms of occurrences of singles, pairs, or triplets of instantaneous phase estimates from the phase time-series (Hillebrand et al., 2016). For our analysis, the number of bins in the histograms was set as 3.49× STD× M −1/3 and delay τ was set as 2M / M±, where STD is average standard deviation of the phase time-series and M ± is the number of times the phase changes sign across time and channels (Hillebrand et al., 2016). PTE has been shown to be robust against the choice of the delayτ and the number of bins for forming the histograms (Hillebrand et al., 2016). In our analysis, PTE was calculated for the entire encoding and recall periods for each trial and then averaged across trials.

Earlier version: Net causal outflow was calculated as the difference between the total outgoing information and total incoming information, that is, net causal outflow = PTE(out) − PTE(in). For example, for calculation of PTE(out) and PTE(in) for the AI electrodes, electrodes in the PCC/precuneus, mPFC, dPPC, and MFG were considered, that is, PTE(out) was calculated as the net PTE from AI electrodes to the PCC/precuneus, mPFC, dPPC, and MFG electrodes, and PTE(in) was calculated as the net PTE from the PCC/precuneus, mPFC, dPPC, and MFG electrodes to AI electrodes. Net causal outflow for the PCC/precuneus, mPFC, dPPC, and MFG electrodes were calculated similarly.

Revised version: Net outflow was calculated as the difference between the total outgoing information and total incoming information, that is, net outflow = PTE(out) − PTE(in). For example, for calculation of PTE(out) and PTE(in) for the AI electrodes, electrodes in the PCC/precuneus, mPFC, dPPC, and MFG were considered, that is, PTE(out) was calculated as the net PTE from AI electrodes to the PCC/precuneus, mPFC, dPPC, and MFG electrodes, and PTE(in) was calculated as the net PTE from the PCC/precuneus, mPFC, dPPC, and MFG electrodes to AI electrodes. Net outflow for the PCC/precuneus, mPFC, dPPC, and MFG electrodes were calculated similarly.

iEEG analysis of phase locking value (PLV) and phase synchronization

Earlier version: We used phase locking value (PLV) to compute phase synchronization between two time-series (Lachaux, Rodriguez, Martinerie, & Varela, 1999). We first calculated the instantaneous phases of the two signals by using the analytical signal approach based on the Hilbert transform (Bruns, 2004). Given time-series x(t), t = 1, 2,…, M, its complex-valued analytical signal z(t) can be computed as where i denotes the square root of minus one, is the Hilbert transform of x(t), and Ax(t) and Φx (t) are the instantaneous amplitude and instantaneous phase respectively and can be given by The Hilbert transform of x(t) was computed as where PV denotes the Cauchy principal value. MATLAB function “hilbert” was used to calculate the Hilbert transform in our analysis. Given two time-series x(t) and y(t), where t = 1, 2,…, M, the PLV (zero-lag) can be computed as where Φy (t) is the instantaneous phase for time-series y(t), denotes the absolute value operator, E [·] denotes the expectation operator with respect to time t, and i denotes the square root of minus one. PLVs were then averaged across trials to estimate the final PLV for each pair of electrodes.

Revised version: We used phase locking value (PLV) to compute phase synchronization between two time-series (Lachaux, Rodriguez, Martinerie, & Varela, 1999). We first calculated the instantaneous phases of the two signals by using the analytical signal approach based on the Hilbert transform (Bruns, 2004). Given time-series x(t), t =1,2,…, M, its complex-valued analytical signal z(t) can be computed as where i denotes the square root of minus one, is the Hilbert transform of x(t), and Axt) and Φx (t) are the instantaneous amplitude and instantaneous phase respectively and can be given by The Hilbert transform of x(t) was computed as where PV denotes the Cauchy principal value. MATLAB function “hilbert” was used to calculate the Hilbert transform in our analysis. Given two time-series x(t) and y(t), where t = 1, 2,…, M, the PLV (zero-lag) can be computed as where Φy(t) is the instantaneous phase for time-series y(t), ∣·∣ denotes the absolute value operator, E [·] denotes the expectation operator with respect to time t, and i denotes the square root of minus one. PLVs were then averaged across trials to estimate the final PLV for each pair of electrodes.

Statistical analysis

Statistical analysis was conducted using mixed effects analysis with the lmerTest package (Kuznetsova, Brockhoff, & Christensen, 2017) implemented in R software (version 4.0.2, R Foundation for Statistical Computing). Because PTE data were not normally distributed, we used BestNormalize (Peterson & Cavanaugh, 2018) which contains a suite of transformation-estimating functions that can be used to optimally normalize data. The resulting normally distributed data were subjected to mixed effects analysis with the following model: PTE ∼ Condition + (1|Subject), where Condition models the fixed effects (condition differences) and (1|Subject) models the random repeated measurements within the same participant, similar to prior iEEG studies (Das & Menon, 2021; Hoy, Steiner, & Knight, 2021; Salamone et al., 2021). Before running the mixed-effects model, PTE was first averaged across trials for each channel pair. Analysis of variance (ANOVA) was used to test the significance of findings with FDR-corrections for multiple comparisons (p<0.05). Linear mixed effects models were run for encoding and recall periods separately. Similar mixed effects statistical analysis procedures were used for comparison of high-gamma power across task conditions, where the mixed effects analysis was run on each of the 0.2s windows.

For effect size estimation, we used Cohen’s d statistics for pairwise comparisons. We used the lme.dscore() function in the EMAtools package in R for estimating Cohen’s d.

Bayesian replication analysis

We used replication Bayes factor (Ly et al., 2019; Verhagen & Wagenmakers, 2014) analysis to estimate the degree of replicability for the direction of information flow for each frequency and task condition and across task domains. Analysis was implemented in R software using the BayesFactor package (Rouder, Speckman, Sun, Morey, & Iverson, 2009). Because PTE data were not normally distributed, as previously, we used BestNormalize (Peterson & Cavanaugh, 2018) to optimally normalize data. We calculated the replication Bayes factor for pairwise experiments. We compared the Bayes factor of the joint model PTE(task1+task2) ∼ Condition + (1|Subject) with the Bayes factor (BF) of individual model as PTE(task1) ∼ Condition + (1|Subject), where task1 denotes the verbal free recall (original) task and task2 denotes the categorized verbal free recall, paired associates learning verbal cued recall, or water maze spatial memory (replication) conditions. We calculated the ratio BF(task1+task2)/BF(task1), which was used to quantify the degree of replicability. We determined whether the degree of replicability was higher than 3 as Bayes factor of at least 3 indicates evidence for replicability (Jeffreys, 1998). A Bayes factor of at least 100 is considered as “decisive” for degree of replication (Jeffreys, 1998). Same analysis procedures were used to estimate the degree of replicability for high-gamma power comparison of DMN and FPN electrodes with the AI electrodes, across experiments.